# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Thomas|secret|
|Agaath|secret|
|Jesse|secret|
|Sam|secret|

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-Kennis-en-Data" data-toc-modified-id="Assignment-Kennis-en-Data-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Assignment Kennis en Data</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Peerreview" data-toc-modified-id="Peerreview-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Peerreview</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.4"><span class="toc-item-num">1.4&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li><li><span><a href="#Na-het-inleveren!" data-toc-modified-id="Na-het-inleveren!-1.5"><span class="toc-item-num">1.5&nbsp;&nbsp;</span>Na het inleveren!</a></span></li></ul></li><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-2.1"><span class="toc-item-num">2.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-2.1.1"><span class="toc-item-num">2.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-2.1.2"><span class="toc-item-num">2.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-2.1.2.1"><span class="toc-item-num">2.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-2.1.2.2"><span class="toc-item-num">2.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-2.1.2.3"><span class="toc-item-num">2.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
import copy
studnr=11873736 # verander door het student id van jullie groepsleider
complete= 'complete_'+str(studnr)
print(complete,"\n")

def complete_11873736(facts, rules):
    """Returnt een lijst met triples van waarheden die zijn geconcludeerd uit de facts en rules."""
    extra_store = facts.copy()
    completion = facts.copy()

    # Loop over de sets in completion
    for triple in completion:
        # Loop over de sets met key subPropertyOf
        for tuples in rules['subPropertyOf']:
            if triple[1] == tuples[0]:
                extra_store.add((triple[0], tuples[1], triple[2]))
        # Loop over de sets met key subClassOf
        for tuples in rules['subClassOf']:
            if triple[2] == tuples[0]:
                extra_store.add((triple[0], triple[1], tuples[1]))

        # Loop over de sets met key domain
        for tuples in rules['domain']:
            if triple[1] == tuples[0]:
                extra_store.add((triple[0], 'isa', tuples[1]))
        # Loop over de sets met key range
        for tuples in rules['range']:
            if triple[1] == tuples[0]:
                extra_store.add((triple[2], 'isa', tuples[1]))

    # Als true dan zijn er geen extra toevoegingen
    if extra_store == completion:
        return completion
    # Voegt nieuwe triples toe aan completion en runt de functie opnieuw uit
    else:
        completion.update(extra_store)
        complete_11873736(completion, rules)
    #return completion
    return completion

#test
facts= test1[0][0]
rules= test1[0][1]
print(facts)
print(rules)
complete_11873736(facts,rules)
```

Output:
```
complete_11873736 

{('Henk', 'isa', 'Mens')}
{'subPropertyOf': set(), 'subClassOf': {('Mens', 'Sterfelijk')}, 'domain': set(), 'range': set()}
{('Henk', 'isa', 'Mens'), ('Henk', 'isa', 'Sterfelijk')}
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-5-919be6bc59f7> in <module>
      4       'range': ('Vindt_lekker','Eten')}
      5 RDF= {('Henk','isa','Mens')}
----> 6 completion= functienaam(RDF,RDFS)
      7 assert isinstance(completion,set)
      8 for t in completion:

NameError: name 'functienaam' is not defined
```

# Jouw lijst met testen

```python
# Checker functie maakt alle mogelijke triples en checkt of de negatieve en positieve voorbeelden kloppen
def checker(test):
    completion = test[0]
    completed = complete_11873736(test[0],test[1])
    positives=test[2]
    negatives=test[3]
    """Print alle mogelijke triples voor de test"""
    print("Alle mogelijke triples voor deze test")
    print(completion)
    """Print true als alle positieve voorbeelden en geen enkele van de negatieve voorbeelden in completion zitten."""
    print("----------------------------------------------------")
    print("Zitten alle positieve voorbeelden in de completion?")
    print(all([True for triple in positives if triple in completed])==True)
    print("----------------------------------------------------")
    print("Zitten er geen negatieve voorbeelden in de completion?")
    print(all([True for triple in negatives if triple not in completed])==True)
    print("----------------------------------------------------")
```

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(studnr)


# Als Henk een mens is dan is hij ook een homo, mensachtige, aap, primaat, zoogdier, choradier en dier.
test2= ({('henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Mens','Homo Sapiens'), ('Homo Sapiens','Mensachtige'), ('Mensachtige','Aap'), ('Aap','Primaat'), ('Primaten','Zoogdier'), ('Zoogdier','Choradier'), ('Choradier','Dier')},
        'domain': set(),
        'range':  set()},
       {('henk','isa','Homo Sapiens'),('henk','isa','Mensachtige'),('henk','isa','Aap'),('henk','isa','Primaat'),('henk','isa','Zoogdier'),('henk','isa','Choradier'),('henk','isa','Dier')},
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )

# Als Henk een mens is dan is hij ook een homo, mensachtige, aap, primaat, zoogdier, choradier en dier (maar deze keer de subclasses achterstevoren).
test3= ({('henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Choradier','Dier'), ('Zoogdier','Choradier'), ('Primaat','Zoogdier'), ('Aap','Primaat'), ('Mensachtige','Aap'), ('Homo Sapiens','Mensachtige'), ('Mens','Homo Sapiens')},
        'domain': set(),
        'range':  set()},
       {('henk','isa','Homo Sapiens'),('henk','isa','Mensachtige'),('henk','isa','Aap'),('henk','isa','Primaat'),('henk','isa','Zoogdier'),('henk','isa','Choradier'),('henk','isa','Dier')},
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )

# Zwitserland is een land met bergen en grotten en het grenst aan Frankrijk. Frankrijk heeft een kust maar Zwitserland niet.
test4= ({('switzerland', 'isa', 'Country'), ('switzerland', 'Contains', 'Mountains'), ('Mountains', 'Possesses', 'Caves'), ('switzerland', 'Neighbours', 'france'), ('france', 'Owns', 'Seashores')},  # RDF store
       {'subPropertyOf': {},  # RDFS store
        'subClassOf': {},
        'domain': {('Neighbours', 'Country')},
        'range':  {('Neighbours', 'Country')}},
       {('switzerland', 'isa', 'Country'), ('switzerland', 'Possesses', 'Caves')}, # positieve voorbeelden
       {('switzerland', 'Owns', 'Seashores'), ('switzerland', 'Neighbours', 'Seashores'), ('france', 'Neighbours', 'Caves')} # negatieve voorbeelden
       )

# Lege sets
test5= (set(),  # RDF store
       {'subPropertyOf': set(),  # RDFS store
        'subClassOf': set(),
        'domain': set(),
        'range':  set()},
       set(), # positieve voorbeelden
       set() # negatieve voorbeelden
       )

# Pieter vind pizza lekker, maar is geen Eten
test6= ({('pieter', 'Vindt_lekker', 'Pizza'), ('Pizza','isa','Eten'), ('Pizza', 'Has', 'toppings'),('Ham', 'isa', 'Topping')},  # RDF store
       {'subPropertyOf':{('Eet','Vindt_lekker')},  # RDFS store
        'subClassOf': {('Pizza','Eten'),('Mens', 'Dier'), ('Ham', 'Topping'),('Ham', 'Eten')},
        'domain': {('Vindt_lekker','Eten')},
        'range':  set()},
       {('Pizza', 'isa', 'Eten'), ('Pieter', 'isa', 'Mens') }, # positieve voorbeelden
       {('pieter', 'isa', 'Eten')} # negatieve voorbeelden
       )

# Jakarta is de hoofdstad van Indonesie, maar Jakarta grenst niet aan een land
test7= ({('Indonesie','isa','Country'), ('Jakarta','isa','Capital'),
       ('Australia','isa','Country'), ('australia','Neighbours','indonesia'),
       ('indonesie','HasCapital','jakarta'), ('indonesie','HasPresident','joko widodo'),
       (('jakarta','HasPopulation',9769000))},  # RDF store
     {'subPropertyOf': {('Isin','HasCapital')},  # RDFS store
      'subClassOf': {('Capital','City'), ('indonesie','Country')},
      'domain': {('Neighbours','Country'), ('Isin','City'), ('HasCapital','Country'), ('HasPresident','Country')},
      'range': {('Neighbours','Country'), ('Isin','Country'), ('HasCapital','Capital'), ('HasPresident','Mens')}},
       {('jakarta','isa','City'), ('jakarta','isin','indonesie')}, # positieve voorbeelden
       {('jakarta','Neighbours','Country')} # negatieve voorbeelden
       )

# Vogels leggen eieren maar Wolven doen dat niet
test8= ({('Owl', 'isa', 'Bird'), ('Bird', 'Lay', 'Eggs'), ('Wolf', 'isa', 'Mammal'), ('Falcon', 'isa', 'Bird')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Bird', 'Animal'), ('Mammal', 'Animal')},
        'domain': set(),
        'range':  set()},
       {('Owl', 'Lay', 'Eggs'), ('Falcon', 'Lay', 'Eggs')}, # positieve voorbeelden
       {('Wolf', 'Lay', 'Eggs')} # negatieve voorbeelden
       )

# Jan houdt van Lotte, maar Bram is getrouwd met Lotte. Als je getrouwd bent ben je een mens, Bram is een mens maar het is niet zeker of Jan ook een mens is.
test9= ({('jan', 'Loves', 'lotte'), ('bram', 'Knows', 'jan'), ('bram', 'Married', 'lotte'), ('bram', 'isa', 'Human'), ('jan', 'leeftijd', '105')},  # RDF store
       {'subPropertyOf': {('Married', 'Loves'), ('Loves', 'Knows')},  # RDFS store
        'subClassOf': {('Human', 'Animal')},
        'domain': {('Married', 'Human')},
        'range':  {('Married', 'Human'), ('Loves', 'Human')}},
       {('bram', 'Knows', 'lotte'), ('lotte', 'isa', 'Human')}, # positieve voorbeelden
       {('jan', 'Married', 'lotte'), ('jan', 'isa', 'Human'), ('jan', 'isa', 'Animal')} # negatieve voorbeelden
       )

# zet hier je tests in
test_11873736 = [test2, test3, test4, test5, test6, test7, test8, test9]

# Vul hier de te checken test in
for t in test_11873736:
    checker(t)

```

Output:
```
Alle mogelijke triples voor deze test
{('henk', 'isa', 'Mens'), ('henk', 'isa', 'Mensachtige'), ('henk', 'isa', 'Homo Sapiens'), ('henk', 'isa', 'Primaat'), ('henk', 'isa', 'Aap')}
----------------------------------------------------
Zitten alle positieve voorbeelden in de completion?
True
----------------------------------------------------
Zitten er geen negatieve voorbeelden in de completion?
True
----------------------------------------------------
Alle mogelijke triples voor deze test
{('henk', 'isa', 'Mens'), ('henk', 'isa', 'Zoogdier'), ('henk', 'isa', 'Mensachtige'), ('henk', 'isa', 'Homo Sapiens'), ('henk', 'isa', 'Primaat'), ('henk', 'isa', 'Dier'), ('henk', 'isa', 'Choradier'), ('henk', 'isa', 'Aap')}
----------------------------------------------------
Zitten alle positieve voorbeelden in de completion?
True
----------------------------------------------------
Zitten er geen negatieve voorbeelden in de completion?
True
----------------------------------------------------
Al
```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.