# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Matthijs| secret |
|Rico| secret |
|Friso| secret |
|Stijn| secret |

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-1.1.1"><span class="toc-item-num">1.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-1.1.2"><span class="toc-item-num">1.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-1.1.2.1"><span class="toc-item-num">1.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-1.1.2.2"><span class="toc-item-num">1.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-1.1.2.3"><span class="toc-item-num">1.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-4.0.1"><span class="toc-item-num">4.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
studnr=12716979 # verander door het student id van jullie groepsleider
functienaam= 'complete_'+str(studnr)
print(functienaam)

def complete_12716979(facts, rules):
    '''Voor elke feit (o,p,s), facts+rules \models⊨ (o,p,s) \Leftrightarrow ⇔ complete(facts, rules) \models⊨ (o,p,s)'''
    for fact in facts.copy():
        
        # split triple in object, predicate and subjcet
        o,p,s = fact[0], fact[1], fact[2]
        
        # loop through the different rules and restrictions
        for rule,restrictions in rules.items():
            for restriction in restrictions:
                
                # create dict with what triple has to be added with every rule
                rule_dict = {
                    "subPropertyOf" :[p, (o, restriction[1], s)],
                    "subClassOf" :  [s, (o, 'isa', restriction[1])],
                    "domain" : [p, (o, 'isa', restriction[1])],
                    "range" : [p, (s, 'isa', restriction[1])]
                }
                
                check = rule_dict[rule][0]
                triple = rule_dict[rule][1]
                
                # check if restriction of current triple is in the current rule
                if restriction[0] == check and triple not in facts:
                    facts.add(triple)
                    complete_12716979(facts, rules)
    return facts
    


#test
facts= test1[0][0]
rules= test1[0][1]
print(facts)
print(rules)
complete_12716979(facts,rules)
```

Output:
```
complete_12716979
{('Henk', 'isa', 'Mens')}
{'subPropertyOf': set(), 'subClassOf': {('Mens', 'Sterfelijk')}, 'domain': set(), 'range': set()}
{('Henk', 'isa', 'Mens'), ('Henk', 'isa', 'Sterfelijk')}
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

```python

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': {('Vindt_lekker','Eten')}
      }
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-5-6fc02bc1cbf3> in <module>
      5       }
      6 RDF= {('Henk','isa','Mens')}
----> 7 completion= functienaam(RDF,RDFS)
      8 assert isinstance(completion,set)
      9 for t in completion:

TypeError: 'str' object is not callable
```

# Jouw lijst met testen

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(studnr)
def test_functie(test):
    facts = test[0]
    rules = test[1]
    completion = complete_12716979(facts, rules)
    print(completion)
    print('Positieve voorbeelden: ', test[2].issubset(completion))
    if not test[2].issubset(completion):
        print("Foute positieve voorbeelden: ", [n for n in test[2] if n not in completion])
    neg = [triple for triple in test[3] if triple in completion]
    print('Geen negatieve voorbeelden: ', len(neg) == 0)
    if not len(neg) == 0:
        print("Foute negatieve voorbeelden: ", neg)


''' We were not allowed to use test2 and test3 according to Maarten.
    But they work on our function so we still wanted to put them here '''

# Empty RDFS test
test2 = (
    {('Klaas', 'isa', 'AOW-trekker'), ('Peter', 'Belt_met', 'Piet')},
    dict(),
    {('Klaas', 'isa', 'AOW-trekker')},     # Positieve voorbeelden
    {('Klaas', 'Belt_met', 'Peter')}    # Negatieve voorbeelden
)

# Missing Domain key
test3 = (
    {('Klaas', 'isa', 'AOW-trekker'), ('Peter', 'Belt_met', 'Piet')},
    {'subPropertyOf': {('Houdt_van', 'Kent'), ('Belt_met', 'Praat_met')},
     'subClassOf': {('Mens', 'Omnivoor'), ('Bejaard', 'Mens'), ('Omnivoor', 'Wezen'), ('AOW-trekker', 'Bejaard')},
     'range': set()
     },
     {('Klaas', 'isa', 'Omnivoor')},     # Positieve voorbeelden
     {('Klaas', 'Belt_met', 'Peter')}    # Negatieve voorbeelden
)

# Recursive test
test4 = (
    {('Klaas', 'isa', 'AOW-trekker'), ('Peter', 'Belt_met', 'Piet')},
    {'subPropertyOf': {('Houdt_van', 'Kent'), ('Belt_met', 'Praat_met')},
     'subClassOf': {('Mens', 'Omnivoor'), ('Bejaard', 'Mens'), ('Omnivoor', 'Wezen'), ('Omnivoor', 'Mens'), ('AOW-trekker', 'Bejaard')},
     'domain': {('Praat_met', 'Mens')},
     'range': set()
     },
     {('Klaas', 'isa', 'Omnivoor')},     # Positieve voorbeelden
     {('Klaas', 'Belt_met', 'Peter')}    # Negatieve voorbeelden
)

# Test with integers
test5 = (
    {(1, 'isa', 'natuurlijk getal'), (2, 'is groter dan', 1)},
    {'subPropertyOf': {('is groter dan', 'niet kleiner dan')},
     'subClassOf': {('natuurlijk getal', 'cijfer')},
     'domain': set(),
     'range': {('is groter dan', 'cijfer')}},
    {(1, 'isa', 'natuurlijk getal')},     # Positieve voorbeelden
    {(2, 'isa', 'natuurlijk getal'), ('1', 'isa', 'natuurlijk getal')}    # Negatieve voorbeelden
)


# Wrong restriction, one item in 'subClassOf' restriction ('Voetballer')
test6 = (
    {('Rico', 'isa', 'Voetballer'), ('Peter', 'Houdt_van', 'Johan')},
    {'subPropertyOf': {('Houdt_van', 'Kent'), ('Belt_met', 'Praat_met')},
     'subClassOf': {('Voetballer', 'Mens'), ('Mens', 'Omnivoor'), ('Omnivoor', 'Wezen'), ('Omnivoor', 'Mens'), ('Voetballer', 'Sporter'), ('Sporter','Sportief'), ('Voetballer')},
     'domain': {('Praat_met', 'Mens')},
     'range': set()
     },
    {('Rico', 'isa', 'Omnivoor'), ('Rico', 'isa', 'Sportief')},     # Positieve voorbeelden
    {('Rico', 'Houdt_van', 'Johan')}    # Negatieve voorbeelden)
)

# Test for every RDFS statement
test7 = (
    {('Rico', 'isa', 'Voetballer'), ('Peter', 'Houdt_van', 'Johan'), ('Rico', 'Vindt_lekker', 'Pizza')},
    {'subPropertyOf': {('Houdt_van', 'Kent'), ('Belt_met', 'Praat_met')},
     'subClassOf': {('Omnivoor', 'Wezen'), ('Omnivoor', 'Mens'), ('Voetballer', 'Sporter'), ('Sporter','Sportief'), ('Voetballer'), ('Voetballer', 'Mens'), ('Mens', 'Omnivoor')},
     'domain': {('Praat_met', 'Mens')},
     'range': {('Vindt_lekker', 'Eten')}
     },
    {('Rico', 'isa', 'Mens'), ('Rico', 'isa', 'Voetballer'), ('Rico', 'isa', 'Sporter'), ('Peter', 'Houdt_van', 'Johan'), ('Peter', 'Kent', 'Johan'), ('Rico', 'isa', 'Wezen'), ('Rico', 'isa', 'Omnivoor'), ('Rico', 'isa', 'Sportief'), ('Pizza', 'isa', 'Eten')},     # Positieve voorbeelden
    {('Rico', 'Houdt_van', 'Johan'), ('Johan', 'isa', 'Mens')}    # Negatieve voorbeelden)
)

# Test with empty RDFS
test8 = (
    {('Klaas', 'isa', 'AOW-trekker'), ('Peter', 'Belt_met', 'Piet')},
    {'subPropertyOf': set(),
     'subClassOf': set(),
     'domain': set(),
     'range': set()
     },
    {('Klaas', 'isa', 'AOW-trekker')},     # Positieve voorbeelden
    {('Klaas', 'Belt_met', 'Peter')}    # Negatieve voorbeelden
)

# Test with floats
test9 = (
    {(1.0, 'isa', 'cijfer'), (20.98, 'is groter dan', 20.1), (21, 'isa', 'Getal'), ('K', 'isa', 'letter')},
    {'subPropertyOf': {('is groter dan', 'niet kleiner dan')},
     'subClassOf': {('cijfer', 'Getal'),('letter', 'karakter')},
     'domain': {('is groter dan', 'cijfer')},
     'range': {('is groter dan', 'cijfer')}
     },
    {(20.98, 'isa', 'cijfer'),('K', 'isa', 'karakter')},     # Positieve voorbeelden
    {(21, 'isa', 'cijfer'), (20.1, 'isa', 'letter')}    # Negatieve voorbeelden
)

# Test only on Domain and Range
test10 = (
    {('Rico', 'vindt_lekker', 'pizza')},
    {'subPropertyOf': set(),
     'subClassOf': set(),
     'domain': {('vindt_lekker', 'Mens')},
     'range': {('vindt_lekker', 'Eten')}
     },
     {('Rico', 'isa', 'Mens')},     # Positieve voorbeelden
     {('pizza', 'isa', 'Mens')}    # Negatieve voorbeelden
)
# Tests for our function
our_tests = [test2, test3, test4, test5, test6, test7, test8, test9, test10]

# Tests for other people
test_12716979= [test4, test5, test6, test7, test8, test9, test10] # zet hier je tests in

# Running tests
for i,test in enumerate(our_tests):
    print(f"Test {i + 2}:")
    test_functie(test)
    print("")
```

Output:
```
Test 2:
{('Peter', 'Belt_met', 'Piet'), ('Klaas', 'isa', 'AOW-trekker')}
Positieve voorbeelden:  True
Geen negatieve voorbeelden:  True

Test 3:
{('Klaas', 'isa', 'Bejaard'), ('Klaas', 'isa', 'AOW-trekker'), ('Klaas', 'isa', 'Mens'), ('Peter', 'Praat_met', 'Piet'), ('Peter', 'Belt_met', 'Piet'), ('Klaas', 'isa', 'Wezen'), ('Klaas', 'isa', 'Omnivoor')}
Positieve voorbeelden:  True
Geen negatieve voorbeelden:  True

Test 4:
{('Klaas', 'isa', 'Bejaard'), ('Klaas', 'isa', 'AOW-trekker'), ('Klaas', 'isa', 'Mens'), ('Peter', 'Praat_met', 'Piet'), ('Peter', 'isa', 'Mens'), ('Peter', 'isa', 'Omnivoor'), ('Peter', 'isa', 'Wezen'), ('Peter', 'Belt_met', 'Piet'), ('Klaas', 'isa', 'Wezen'), ('Klaas', 'isa', 'Omnivoor')}
Positieve voorbeelden:  True
Geen negatieve voorbeelden:  True

Test 5:
{(2, 'niet kleiner dan', 1), (1, 'isa', 'natuurlijk getal'), (2, 'is groter dan', 1), (1, 'isa', 'cijfer')}
Positieve voorbeelden:  True
Geen negatieve voorbeelden:  True

Test 6:
{('Rico', 'isa', 'Mens'), ('Ri
```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.

```python
def complete(RDFS):
    
    # Loop through the different rules and restrictions
    for rule, restrictions in RDFS.items():

        # Domain influences subClassOf, subPropertyOf influences range
        rule_dict = {
            "subPropertyOf" : restrictions.copy(),
            "subClassOf" : restrictions.copy(),
            "domain" : RDFS['subClassOf'],
            "range" : RDFS['subPropertyOf']
        }
        
        # Loop through the current restrictions and the correct dependent restrictions
        for (o, s) in restrictions.copy():
            for (o2, s2) in rule_dict[rule]:
                
                # Create the right checks for the right rules
                check_dict = {
                    "subPropertyOf" : [(s == o2), (o, s2)],
                    "subClassOf" : [(s == o2), (o, s2)],
                    "domain" : [(s == o2), (o, s2)],
                    "range" : [(o == s2), (o2, s)]
                }
                checks = check_dict[rule]
                
                # make sure restriction is not yet in RDFS list
                if checks[0] and checks[1] not in RDFS[rule]:
                    RDFS[rule].add(checks[1])
                    complete(RDFS)
    return RDFS

# Test function
def test_functie2(test):
    completion = complete(test[0])
    print('\n', completion, '\n')
    for rule, values in completion.items():
        pos_list = [p for p in test[1][rule] if p not in values]
        print(f'All positive examples {rule}: ', len(pos_list) == 0)
        neg_list = [p for p in test[2][rule] if p in values]
        print(f'No negative examples {rule}: ', len(neg_list) == 0)
        print('')


# Test lists in format rdfs, positives, negatives
test_rdfs1 = (
    {
        'subPropertyOf': {('Houdt_van','Is_vrienden_met'), ('Is_vrienden_met', 'Kent'), ('Eet','Vindt_lekker')},
        'subClassOf': {('Mens','Dier'),('Student','Mens')},
        'domain': {('Houdt_van','Mens')},
        'range': {('Vindt_lekker','Eten')}
    },
    {
        'subPropertyOf': {('Houdt_van', 'Kent')},
        'subClassOf': set(),
        'domain': {('Houdt_van', 'Mens')},
        'range': {('Eet', 'Eten')}
    },
    {
        'subPropertyOf': {('Is_vrienden_ment', 'Houdt_van')},
        'subClassOf': set(),
        'domain': ('Houdt_van', 'Student'),
        'range': {'Vindt_lekker', 'Dier'}
    }
)

test_rdfs2 = (
    {
    'subPropertyOf': {('Houdt_van', 'Kent'), ('Eet', 'Vindt_lekker')},
    'subClassOf': {('Mens', 'Omnivoor'), ('Omnivoor', 'Wezen')},
    'domain': {('Praat_met', 'Mens'), ('Praat_met', 'Wezen'), ('Vindt_lekker', 'Mens')},
    'range': {('Vindt_lekker', 'Eten')}
    },
    {
        'subPropertyOf': {('Houdt_van', 'Kent')},
        'subClassOf': {('Mens', 'Wezen')},
        'domain': {('Praat_met', 'Omnivoor')},
        'range': {('Eet', 'Eten')}
    },
    {
        'subPropertyOf': {('Is_vrienden_met', 'Houdt_van')},
        'subClassOf': set(),
        'domain': set(),
        'range': {'Vindt_lekker', 'Dier'}
    }
)

rdfs_tests = [test_rdfs1, test_rdfs2]

for i,test in enumerate(rdfs_tests):
    print(f"Test {i + 1}:")
    test_functie2(test)
    print("")

```

Output:
```
Test 1:

 {'subPropertyOf': {('Houdt_van', 'Kent'), ('Eet', 'Vindt_lekker'), ('Is_vrienden_met', 'Kent'), ('Houdt_van', 'Is_vrienden_met')}, 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens'), ('Student', 'Dier')}, 'domain': {('Houdt_van', 'Mens'), ('Houdt_van', 'Dier')}, 'range': {('Eet', 'Eten'), ('Vindt_lekker', 'Eten')}} 

All positive examples subPropertyOf:  True
No negative examples subPropertyOf:  True

All positive examples subClassOf:  True
No negative examples subClassOf:  True

All positive examples domain:  True
No negative examples domain:  True

All positive examples range:  True
No negative examples range:  True


Test 2:

 {'subPropertyOf': {('Houdt_van', 'Kent'), ('Eet', 'Vindt_lekker')}, 'subClassOf': {('Mens', 'Omnivoor'), ('Omnivoor', 'Wezen'), ('Mens', 'Wezen')}, 'domain': {('Vindt_lekker', 'Wezen'), ('Praat_met', 'Mens'), ('Vindt_lekker', 'Omnivoor'), ('Praat_met', 'Wezen'), ('Vindt_lekker', 'Mens'), ('Praat_met', 'Omnivoor')}, 'range': {('Eet', 'Eten'), ('Vind
```

# Uitleg code
De gemaakte code voor het aanvullen van de RFDS statements is correct omdat de juiste output wordt weergegeven. De code vult de juiste "subPropertyOf" aan. Zo wordt de tuple ('houdt_van', 'kent') toegevoegd in test1 aan de hand van de tuples ('houdt_van', 'is_vrienden_met') & ('is_vrienden_met', 'kent'). Bij 'subClassOf' wordt in test 1 ook juist de tuple ('Student', 'Dier') aan de hand van de gegeven tuples ('Student', 'Mens') & ('Mens', 'Dier'). "subPropertyOf" en "subClassOf" hebben dus direct invloed op zichzelf, "domain" en "range" werken daarentegen op een iets andere manier. "subClassOf" heeft namelijk invloed op "domain", want zoals je ziet in test1, als je de subclass restricties ("Mens", "Dier") & ("Student", "Mens") en de domain restrictie ("Houdt_van", "Mens"), dan kan ("Houdt_van", "Dier") als domain restrictie toegevoegd worden, omdat Mens dus subclass is van Dier. Andersom kan ("Houdt_van", "Student") niet toegevoegd worden als domain restrictie, omdat deze niet direct afgeleid kan worden uit de restrictie ("Houdt_van", "Mens"). Als x van y houdt is x een mens, maar hoeft x dus niet per se een student te zijn, maar wel zeker een dier, omdat mens subclass is van mens. Hiernaast heeft "subPropertyOf" invloed op "range". Laten we weer kijken naar test1, deze test heeft de "subPropertyOf" restrictie ("Eet", "Vindt_lekker"), dus als je de triple (x, "Eet", pizza) hebt, dan kan je afleiden (x, "Vindt_lekker", pizza). Test1 heeft ook de "range" restrictie ("Vindt_lekker", "Eten"), dus als (x, "Vindt_lekker", pizza) dan (pizza, "isa", "Eten"). Uit deze twee restricties kan dan ook de "range" restrictie ("Eet", "Eten") afgeleid worden, dus zo kan je direct uit de triple (x, "Eet", pizza) de triple (pizza "isa" "Eten") afleiden.

```python

```