# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Robert|secret|
|Nathan|secret|
|Suzanne|secret|
|Berkay|secret|

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-Kennis-en-Data" data-toc-modified-id="Assignment-Kennis-en-Data-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Assignment Kennis en Data</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Peerreview" data-toc-modified-id="Peerreview-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Peerreview</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.4"><span class="toc-item-num">1.4&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li><li><span><a href="#Na-het-inleveren!" data-toc-modified-id="Na-het-inleveren!-1.5"><span class="toc-item-num">1.5&nbsp;&nbsp;</span>Na het inleveren!</a></span></li></ul></li><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-2.1"><span class="toc-item-num">2.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-2.1.1"><span class="toc-item-num">2.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-2.1.2"><span class="toc-item-num">2.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-2.1.2.1"><span class="toc-item-num">2.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-2.1.2.2"><span class="toc-item-num">2.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-2.1.2.3"><span class="toc-item-num">2.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
studnr=11305320 # verander door het student id van jullie groepsleider
functienaam= 'complete_' + str(studnr)
print(functienaam)

def complete_11305320(facts, rules):  
    completion = facts
    
    while True:
        storeSize = len(completion)
        loopstore = {triple for triple in completion}
        
        for triple in loopstore:
            for subProp in rules['subPropertyOf']:
                if triple[1] == subProp[0]:
                    completion.add((triple[0], subProp[1], triple[2]))
                    
            for subClass in rules['subClassOf']:
                if triple[2] == subClass[0]:
                    completion.add((triple[0], triple[1], subClass[1]))
                    
            for dom in rules['domain']:
                if triple[1] == dom[0]:
                    completion.add((triple[0], 'isa', dom[1]))
                    
            for ran in rules['range']:
                if triple[1] == ran[0]:
                    completion.add((triple[2], 'isa', ran[1]))

        if storeSize == len(completion):
            break
            
    return completion

#test

functienaam = complete_11305320

complete_11305320(test1[0][0],test1[0][1])
```

Output:
```
complete_11305320
{('Henk', 'isa', 'Mens'), ('Henk', 'isa', 'Sterfelijk')}
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

# Jouw lijst met testen

```python
from pprint import pprint

def testchecker(tests, print_completion=False):
    for n, test in enumerate(tests):
        print(f"Test #{n}")
        completion = complete_11305320(test[0],test[1])
        for pos in test[2]:
            if pos not in completion:
                print(f"{pos} is not in completion, but should be!")
            else:
                print(f"{pos} correctly in completion")
        for neg in test[3]:
            if neg in completion:
                print(f"{neg} is in completion, but shouldn't be!")    
            else:
                print(f"{neg} correctly not in completion")
        if print_completion:
            print("---")
            pprint(completion)
            print("---")
        print(f"#{n} Getest!\n---")
```

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(studnr)
# jouw test

test_11305320 = [
    ({('irma','isa','Mens'), ('marie','isa','Levend')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Levend','Sterfelijk'),('Sterfelijk', 'Levend'),('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('irma','isa','Sterfelijk'),('irma','isa','Levend'),('marie','isa','Sterfelijk')}, # positieve voorbeelden
       {('marie', 'isa', 'Mens')} # negatieve voorbeelden
        ),
    
    ({('ash ketchum','Vangt','pikachu'),('misty','Vangt','psyduck'),('psyduck','Vangt','pikachu')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Pokémon','Fabieldier'),('Trainer','Mens')},
      'domain': {('Vangt','Trainer')},
      'range':  {('Vangt','Pokémon')}},
       {('pikachu','isa','Fabieldier'),('misty','isa','Mens'),('psyduck','isa','Trainer'),('psyduck','isa','Pokémon')}, # positieve voorbeelden
       {('pikachu', 'isa', 'Mens'),('misty','isa','Pokémon')} # negatieve voorbeelden
       ),
    
    ({('james','Geboorteplaats','londen'), ('nathan', 'Geboorteplaats', 'amsterdam')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': set(),
      'range': {('Geboorteplaats', 'Plaats')}},
       {('londen','isa','Plaats'), ('amsterdam', 'isa', 'Plaats')}, # positieve voorbeelden
       {('james', 'isa', 'Plaats'), ('james', 'Geboorteplaats', 'nathan'), ('londen', 'Geboorteplaats', 'amsterdam')} # negatieve voorbeelden
       ),
    
    ({('cato', 'Legt_neer', 'boek'), ('peter', 'Gooit', 'tafel')},  # RDF store
       {'subPropertyOf': {('Legt_neer', 'Verplaatst'), ('Gooit', 'Verplaatst')},  # RDFS store
      'subClassOf': set(),
      'domain': set(),
      'range': {('Verplaatst', 'Voorwerp')}},
       {('boek','isa','Voorwerp'), ('cato', 'Verplaatst', 'boek'), ('tafel', 'isa', 'Voorwerp')}, # positieve voorbeelden
       {('boek', 'Gooit', 'peter'), ('tafel', 'Legt_neer', 'boek')} # negatieve voorbeelden
       ),
    
     ({('pizza','Contains','cheese'), ('cheese','Contains','salt')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': {('Contains', 'Junkfood')},
      'range': set()},
       {('pizza','isa','Junkfood'), ('cheese','isa','Junkfood')}, # positieve voorbeelden
       {('cheese', 'isa','Pizza'), ('tomato', 'isa','Junkfood')} # negatieve voorbeelden
       ),
    
     ({('oma','Geboortejaar',1923), ('oma','Geboortedatum',19231025), ('mario', 'Geboortejaar', 1985), ('yoshi', 'isa', 'Dinosaurus')},  # RDF store
       {'subPropertyOf': {('Geboortejaar', 'Geboortedatum')},  # RDFS store
      'subClassOf': set(),
      'domain': {('Geboortedatum', 'Mens')},
      'range': set()},
       {('oma','isa','Mens'), ('mario', 'isa','Mens'), ('oma', 'Geboortedatum', 1923)}, # positieve voorbeelden
       {('Geboortedatum', 'isa','Geboortejaar'), ('yoshi','isa','Mens')} # negatieve voorbeelden
       ),
    
    ({('joost', 'Pint', 50),('lisa', 'Betaalt', 100)},  # RDF store
       {'subPropertyOf': {('Pint', 'Betaalt'),('Betaalt', 'Spendeert')},  # RDFS store
      'subClassOf': set(),
      'domain': set(),
      'range':  set()
       },
       {('joost','Betaalt',50),('joost','Spendeert',50),('lisa','Spendeert',100)}, # positieve voorbeelden
       {('joost','Betaalt',100), ('lisa','Pint',100), ('lisa','Spendeert',50)} # negatieve voorbeelden
       ),
    
    ({('andré','Bakt','taart'),('maythe','Kookt','ei')},  # RDF store
       {'subPropertyOf': {('Kookt', 'Bereidt'), ('Bakt', 'Bereidt')},  # RDFS store
      'subClassOf': {('Eten','Levensmiddel'),('Mens','Wezen')},
      'domain': {('Bereidt','Mens')},
      'range':  {('Bereidt','Eten')}},
       {('ei','isa','Levensmiddel'),('andré','isa','Wezen'),('maythe','Bereidt','ei'),('taart','isa','Eten')}, # positieve voorbeelden
       {('taart', 'isa', 'Wezen'),('andré','Bereidt','maythe')} # negatieve voorbeelden
    )
    ] # zet hier je tests in

# testchecker(test_11305320)

testnaam = test_11305320
```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.

### RDFS chase

Naïeve versie:

Met $r$ relaties en $p$ eigenschappen kan je $r\cdot(r−1)$ veel `subPropertyOf` relaties maken

en dan nog $p\cdot(p−1)$ veel `subClassOf` relaties

en dan nog $2*r*p$ veel `domain` en `range` relaties samen.

---

Chase versie:
Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` je daarvan kan afleiden hangt af van de "regels" die je op de statements toepast, en  overlap met andere `range`, `domain` en `subClassOf` statements. Hoeveel je uiteindelijk uit die éne `subPropertyOf` hebt kunnen afleiden weet je pas aan het eind. Als je reflexiviteit toepast kan je iniedergeval ook gelijk afleiden dat het antecedent ook z'n eigen consequent in een nieuw statement is.

Het systeem is geldig en volledig omdat het letterlijk de regeltjes van het artikel in code bevat, en blijft loopen tot er niks meer afgeleid kan worden.

```python
from copy import deepcopy

def complete(RDFS):
    while True:
        storeSize = sum(len(RDFS[x]) for x in RDFS)
        loopstore = deepcopy(RDFS)
        
        for subPropA in loopstore['subPropertyOf']:
            # Regel 2b: subproperty transitiviteit: als A subproperty van B, dan A ook subproperty waar B subproperty van is
            for subPropB in loopstore['subPropertyOf']:
                if subPropA[1] == subPropB[0]:
                    RDFS['subPropertyOf'].add((subPropA[0], subPropB[1]))
                    
            # Regel 3b: Als A subprop van B is, neem domein waar B links staat over met A ipv B
            for dom in loopstore['domain']:
                if subPropA[1] == dom[0]:
                    RDFS['domain'].add((subPropA[0], dom[1]))
            
            # Regel 4b: Als A subprop van B is, neem range waar B links staat over met A ipv B
            for ran in loopstore['range']:
                if subPropA[1] == ran[0]:
                    RDFS['range'].add((subPropA[0], ran[1]))
                    
            # Regel 6b: Alles wat een subprop-relatie heeft is een subprop van zichzelf        
            RDFS['subPropertyOf'].add((subPropA[0], subPropA[0]))
            RDFS['subPropertyOf'].add((subPropA[1], subPropA[1]))

        for subClassA in loopstore['subClassOf']:
            # Regel 1b: subclass transitiviteit: als A subclass van B, dan A ook subclass waar B subclass van is
            for subClassB in loopstore['subClassOf']:
                if subClassA[1] == subClassB[0]:
                    RDFS['subClassOf'].add((subClassA[0], subClassB[1]))

            # Regel 3c: Als A domein van B is, neem domein van superclassen van B over        
            for dom in loopstore['domain']:
                if dom[1] == subClassA[0]:
                    RDFS['domain'].add((dom[0], subClassA[1]))
            
            # Regel 4c: Als A range van B is, neem range van superclassen van B over
            for ran in loopstore['range']:
                if ran[1] == subClassA[0]:
                    RDFS['range'].add((ran[0], subClassA[1]))
            
            # Regel 5a: Alles wat een subclass-relatie heeft is een subclass van zichzelf
            RDFS['subClassOf'].add((subClassA[0], subClassA[0]))
            RDFS['subClassOf'].add((subClassA[1], subClassA[1]))
            
        for ran in loopstore['range']:
            # Regel 5b: Ook alle classes in een range zijn subclass van zichzelf
            RDFS['subClassOf'].add((ran[1], ran[1]))
            
            # Regel 6d: Ook alle properties in een range zijn subprop van zichzelf
            RDFS['subPropertyOf'].add((ran[0], ran[0]))
            
        for dom in loopstore['domain']:
            # Regel 5b: Ook alle classes in een domein zijn subclass van zichzelf
            RDFS['subClassOf'].add((dom[1], dom[1]))
            
            # Regel 6d: Ook alle properties in een domein zijn subprop van zichzelf
            RDFS['subPropertyOf'].add((dom[0], dom[0]))
            
            # Regel 7: Extensional
            for subClassA in loopstore['subClassOf']:
                for subPropA in loopstore['subPropertyOf']:
                    if subPropA[0] == "isa" and subPropA[1] == dom[0] and subClassA[0] == subClassA[1]:
                        RDFS['subClassOf'].add((subClassA[0], dom[1]))
                
            
        if storeSize == sum(len(RDFS[x]) for x in RDFS):
            break
            
    return RDFS
```

```python
RDFS1 = {'subPropertyOf': {("birthCity","birthPlace")},  
         'subClassOf': {("PopulatedPlace","Place"),("City","PopulatedPlace"), ("Historical Place", "Place")},
         'domain': {("birthPlace", "Person")},
         'range':  {("birthPlace","Place")}}

complete(RDFS1)
```

Output:
```
{'subPropertyOf': {('birthCity', 'birthCity'),
  ('birthCity', 'birthPlace'),
  ('birthPlace', 'birthPlace')},
 'subClassOf': {('City', 'City'),
  ('City', 'Place'),
  ('City', 'PopulatedPlace'),
  ('Historical Place', 'Historical Place'),
  ('Historical Place', 'Place'),
  ('Person', 'Person'),
  ('Place', 'Place'),
  ('PopulatedPlace', 'Place'),
  ('PopulatedPlace', 'PopulatedPlace')},
 'domain': {('birthCity', 'Person'), ('birthPlace', 'Person')},
 'range': {('birthCity', 'Place'), ('birthPlace', 'Place')}}
```

```python
def testchecker_RDFS(tests, print_completion=False):
    for n, test in enumerate(tests):
        print(f"Test #{n}")
        completion = complete(test[0])
        for key in test[1]:
            for pos in test[1][key]:
                if pos not in completion[key]:
                    print(f"{pos} is not in completion, but should be!")
                else:
                    print(f"{pos} correctly in completion")
        for key in test[2]:
            for neg in test[2][key]:
                if neg in completion[key]:
                    print(f"{neg} is in completion, but shouldn't be!")
                else:
                    print(f"{neg} correctly not in completion")
        if print_completion:
            print("---")
            pprint(completion)
            print("---")
        print(f"#{n} Getest!\n---")
```

```python
test_RDFS = [
    ({'subPropertyOf': {("Slow bakt","Bakt"),('Kookt', 'Bereidt'), ('Bakt', 'Bereidt')},  # RDFS store
      'subClassOf': {("Snoepgoed","Eten"),('Eten','Levensmiddel'),('Mens','Wezen')},
      'domain': {('Bereidt','Mens')},
      'range':  {('Bereidt','Eten')}},
     
       {'subPropertyOf': {('Slow bakt', 'Bereidt')},  # RDFS store
      'subClassOf': {('Snoepgoed', 'Levensmiddel')},
      'domain': {('Slow bakt', 'Wezen')},
      'range':  {('Kookt', 'Eten')}}, # positieve voorbeelden
     
       {'subPropertyOf': {('Bakt','Kookt')},  # RDFS store
      'subClassOf': {('Eten','Wezen')},
      'domain': {('Bereidt', 'Eten')},
      'range':  {('Bereidt', 'Mens')}} # negatieve voorbeelden
       ),
    
    ({'subPropertyOf': set(),  # RDFS store
      'subClassOf': {("Person","Agent"),("Agent","Resource")},
      'domain': set(),
      'range':  {("Author","Person")}},
     
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Person', 'Resource')},
      'domain': set(),
      'range':  {("Author","Agent")}}, # positieve voorbeelden
     
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': set(),
      'range':  {("Person","Author")}} # negatieve voorbeelden
       ),
    
    ({'subPropertyOf': {("Twerken","Dansen"),("isa","Viben")},  # RDFS store
      'subClassOf': {("Squid","Cephalopod")},
      'domain': {("Viben","Inkling")},
      'range':  {("Battles","Octopus")}},
     
       {'subPropertyOf': {('Battles', 'Battles'), ("Dansen", "Dansen")},  # RDFS store
      'subClassOf': {('Octopus', 'Inkling'), ("Squid","Squid"), ("Octopus","Octopus")},
      'domain': {('isa', 'Inkling')},
      'range':  {('Battles', 'Inkling')}}, # positieve voorbeelden
     
       {'subPropertyOf': {('isa', 'Dansen')},  # RDFS store
      'subClassOf': {('Octopus', 'Cephalapod')},
      'domain': set(),
      'range':  set()} # negatieve voorbeelden
       ),
    
    ({'subPropertyOf': set(),  # RDFS store
      'subClassOf': {("Color","Rainbow"),("Rainbow","Nature")},
      'domain': {("hasColor", "Color")},
      'range':  set()},
     
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Color', 'Nature')},
      'domain': {('hasColor', 'Rainbow')},
      'range': set()}, # positieve voorbeelden
     
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': {('hasColor', 'hasColor')},
      'range': set()} # negatieve voorbeelden
       ),
    
    ({'subPropertyOf': {("In verhaal","Komt voor in")},  # RDFS store
      'subClassOf': {("Disneyfiguur", "Figuur"),("Duck-familielid", "Eend")},
      'domain': {("Komt voor in", "Disneyfiguur")},
      'range':  set()},
     
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Figuur', 'Figuur')},
      'domain': {('Komt voor in', 'Disneyfiguur'), ('In verhaal', 'Figuur')},
      'range':  set()}, # positieve voorbeelden
     
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': {("Komt voor in", "Eend")},
      'range':  set()} # negatieve voorbeelden
       )
       ]

testchecker_RDFS(test_RDFS)
```

Output:
```
Test #0
('Slow bakt', 'Bereidt') correctly in completion
('Snoepgoed', 'Levensmiddel') correctly in completion
('Slow bakt', 'Wezen') correctly in completion
('Kookt', 'Eten') correctly in completion
('Bakt', 'Kookt') correctly not in completion
('Eten', 'Wezen') correctly not in completion
('Bereidt', 'Eten') correctly not in completion
('Bereidt', 'Mens') correctly not in completion
#0 Getest!
---
Test #1
('Person', 'Resource') correctly in completion
('Author', 'Agent') correctly in completion
('Person', 'Author') correctly not in completion
#1 Getest!
---
Test #2
('Battles', 'Battles') correctly in completion
('Dansen', 'Dansen') correctly in completion
('Octopus', 'Octopus') correctly in completion
('Squid', 'Squid') correctly in completion
('Octopus', 'Inkling') correctly in completion
('isa', 'Inkling') correctly in completion
('Battles', 'Inkling') correctly in completion
('isa', 'Dansen') correctly not in completion
('Octopus', 'Cephalapod') correctly not in completion
#2 Gete
```