# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Luuk|secret |
|Rogier|secret |
|Soufiane|secret |
|Milou|secret
 |

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-Kennis-en-Data" data-toc-modified-id="Assignment-Kennis-en-Data-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Assignment Kennis en Data</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Peerreview" data-toc-modified-id="Peerreview-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Peerreview</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.4"><span class="toc-item-num">1.4&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li><li><span><a href="#Na-het-inleveren!" data-toc-modified-id="Na-het-inleveren!-1.5"><span class="toc-item-num">1.5&nbsp;&nbsp;</span>Na het inleveren!</a></span></li></ul></li><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-2.1"><span class="toc-item-num">2.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-2.1.1"><span class="toc-item-num">2.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-2.1.2"><span class="toc-item-num">2.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-2.1.2.1"><span class="toc-item-num">2.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-2.1.2.2"><span class="toc-item-num">2.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-2.1.2.3"><span class="toc-item-num">2.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
studnr=12806714 # verander door het student id van jullie groepsleider
functienaam= 'complete_'+str(studnr)

def complete_12806714(facts, rules):  
    completion=facts
    objectlist=list(facts)
    objectlistNew= []
#   check of er nieuwe conclusies bij zijn gekomen.
    while len(objectlist)!=len(objectlistNew):
        objectlistNew = objectlist
#       per key van de rdfs store, kijken of iets aan de regels voldoet en op basis daarvan een nieuwe triple maken.
        for i in objectlist:
            for g in rules["subPropertyOf"]:
                if g[0]== i[1]:
                    a =(i[0],g[1],i[2])
                    if a not in objectlist:
                        objectlistNew.append((i[0],g[1],i[2]))
            for g in rules['subClassOf']:
                if g[0] == i[2]:
                     if (i[0],i[1],g[1]) not in objectlist:
                        objectlistNew.append((i[0],i[1],g[1]))   
            for g in rules['domain']:
                if g[0] == i[1]:
                    if (i[0],'isa',g[1]) not in objectlist:
                        objectlistNew.append((i[0],'isa',g[1]))
            for g in rules['range']:
                if g[0] == i[1]:
                    if (i[2], 'isa',g[1]) not in objectlist:
                        objectlistNew.append((i[2], 'isa',g[1]))
            check = len(set(objectlist))
#     returnen in set voor het filteren van dubbele conclusies, en alle facts die meegegeven zijn moeten niet gereturnt worden
    return(set(set(objectlist) - facts))


#test
complete_12806714(test1[0][0],test1[0][1])
```

Output:
```
{('Henk', 'isa', 'Sterfelijk')}
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-13-f784c402e67c> in <module>
      4       'range': ('Vindt_lekker','Eten')}
      5 RDF= {('Henk','isa','Mens')}
----> 6 completion= functienaam(RDF,RDFS)
      7 assert isinstance(completion,set)
      8 for t in completion:

TypeError: 'str' object is not callable
```

# Jouw lijst met testen

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(studnr)

# Bij no_facts testen dat de facts die meegegeven zijn, niet worden gereturnt als conclusie
no_facts = [
    ({('henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Dier')},
      'domain': set(),
      'range':  set()},
       {('henk','isa','Dier')}, # positieve voorbeelden
       {('henk', 'isa', 'Mens')} # negatieve voorbeelden
       )
       ]
# Bij more_facts testen dat de functie werkt bij meer dan één meegegeven fact.
more_facts = [
    ({('henk','isa','Mens'), ('henk','Vindt_lekker','Pizza')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Dier')},
      'domain': set(),
      'range':  {('Vindt_lekker','Eten')}},
       {('henk', 'isa', 'Dier'), ('Pizza', 'isa', 'Eten')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'Pizza')} # negatieve voorbeelden
       )
       ]

# bij circulair testen we of de functie niet oneindig door gaan, dus als henk mens is -> henk dier en niet weer henk -> mens
circulair= [
    ({('henk','isa','Mens'),("hans",'isa','Dier')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Student','Mens'),('Mens','Dier')},
      'domain': set(),
      'range':  set()},
       {('henk','isa','Dier')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
# Bij circulair2 testen wij hetzelfde als circulair maar dan met 3 verschillende subclasses in plaats van 2.
circulair2 = [
    ({('henk','isa','Mens'),("hans",'isa','Dier')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Student','Mens'),('Mens','Dier'), ('Dier','Student')},
      'domain': set(),
      'range':  set()},
       {('hans', 'isa', 'Mens'), ('henk', 'isa', 'Student'), ('hans', 'isa', 'Student'), ('henk', 'isa', 'Dier')}, # positieve voorbeelden
       {('henk', 'isa', 'Mens'),('hans', 'isa','Dier')} # negatieve voorbeelden
       )
       ]
# Bij continuiteit testen dat de conclusies die gemaakt zijn door functie ook weer door de functie gehaald worden. 
# Dus als henk van Ingrid houdt komt er als conclusie dat henk een mens is, Dan zou daaruit weer geconcludeerd moeten worden da
# dat henk een dier is. 
continuiteit = [
    ({('henk','Houdt_van','ingrid')},  # RDF store
       {'subPropertyOf': {('Houdt_van','Kent')},  # RDFS store
      'subClassOf': {('Mens','Dier')},
      'domain': {('Houdt_van','Mens')},
      'range':  {('Houdt_van','Mens')}},
       {('henk', 'isa', 'Mens'), ('ingrid', 'isa', 'Mens'), ('ingrid', 'isa', 'Dier'), ('henk', 'Kent', 'Ingrid'), ('henk', 'isa', 'Dier')}, # positieve voorbeelden
       {('ingrid', 'Houdt_van', 'henk')} # negatieve voorbeelden
       )
       ]
# Bij no_conclusion testen dat wanneer er geen facts voldoen aan de regels, er geen conclusies worden gereturnt.
no_conclusion= [
    ({('henk','Houdt_van','ingrid')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Dier')},
      'domain': set(),
      'range':  set()},
       set(), # positieve voorbeelden
       {('henk', 'isa', 'Dier')} # negatieve voorbeelden
       )
       ]

test_12806714=[(no_facts[0][0],no_facts[0][1],no_facts[0][2],no_facts[0][3])
          ,(more_facts[0][0],more_facts[0][1],more_facts[0][2],more_facts[0][3])
          ,(circulair[0][0],circulair[0][1], circulair[0][2],circulair[0][3])
          ,(circulair2[0][0],circulair2[0][1], circulair2[0][2],circulair2[0][3])
          ,(continuiteit[0][0],continuiteit[0][1], continuiteit[0][2], continuiteit[0][3]),
          (no_conclusion[0][0],no_conclusion[0][1],no_conclusion[0][2],no_conclusion[0][3])] # zet hier je tests in


# raise NotImplementedError()

```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-16-500b0aa181b3> in <module>
----> 1 assert isinstance(testnaam,list)
      2 for test in testnaam:
      3     assert isinstance(test,tuple)
      4     assert len(test)==4
      5     assert isinstance(test[0],set)

AssertionError:
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.

```python
def complete(RDFS):
#     counter
    n=0
    RDFS_list= [j for i in RDFS.values() for j in i]
    RDFS_dict = {'subPropertyOf': [], 'subClassOf':[], 'domain':[], 'range':[]}
#   voegt bestaande regels toe aan nieuwe dict.
    for RDF in RDFS:
        for i in RDFS[RDF]:
            RDFS_dict[RDF].append(i)
    for RDF in RDFS:
        for RDF2 in RDFS: 
            for i in RDFS[RDF]:
                for j in RDFS[RDF2]:
#                    checks of de tweede value van i gelijk is aan eerste value van j, dan moet
#                    de eerste waarde van i en tweede waarde van j toegevoegd worden als nieuw regel.
                    if i!=j and i[1]==j[0] and i[0]!=j[1] and (i[0],j[1]) not in [j for i in RDFS.values() for j in i]:
                        if RDF2 in ['range','domain'] and RDF != RDF2:
                            RDFS_dict[RDF2].append((i[0],j[1]))
                            n+=1
                        else:
                            RDFS_dict[RDF].append((i[0],j[1]))
                            n+=1
#                    checks of een waarde van i overeenkomt met een waarde van j. Als deze dan ook in range of domain zit dan
#                    wordt deze dan ook aan die categorie toegevoegd
                    if i!=j and RDF in ['range','domain'] and RDF != RDF2: 
                        if i[1]==j[1] and (i[0],j[0]) not in [j for i in RDFS.values() for j in i]:
                            RDFS_dict[RDF].append((i[0],j[0]))
                            n+=1
                        elif i[0]==j[0] and (j[1],i[1]) not in [j for i in RDFS.values() for j in i]:
                            RDFS_dict[RDF].append((j[1],i[1]))
                            n+=1
#     Als er nieuwe waardes zijn toegevoegd dan roept de functie zichzelf aan met de nieuwe dict, totdat er niks nieuws wordt toegevoegd
    if n!=0:
        complete(RDFS_dict)
    RDFS_dict = {key: set(value) for key,value in RDFS_dict.items()}
    return(RDFS_dict)

RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens'),('Dier','Mens'),('Mens','Student')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}  
# test of de functie werkt. Er wordt gestest op fouten in de output
rdf_test= [(
       {'subPropertyOf':  {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},  # RDFS store
      'subClassOf': {('Mens','Dier'),('Student','Mens'),('Dier','Mens'),('Mens','Student')},
      'domain': {('Houdt_van','Mens')},
      'range':  {('Vindt_lekker','Eten')}},
       {'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')}, 
        'subClassOf': {('Student', 'Dier'), ('Dier', 'Student'), ('Student', 'Mens'), ('Mens', 'Student'), ('Dier', 'Mens'), ('Mens', 'Dier')}, 
        'domain': {('Houdt_van', 'Mens'), ('Houdt_van', 'Dier'), ('Houdt_van', 'Student'), ('Kent', 'Mens')}, 
        'range': {'Eten', 'Vindt_lekker'}}, # positieve voorbeelden
       {'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')}, 
        'subClassOf': {('Student', 'Dier'), ('Mens', 'Student'), ('Houdt_van', 'Dier'), ('Student', 'Mens'), ('Mens', 'Student'), ('Dier', 'Mens'), ('Mens', 'Dier')}, 
        'domain': {('Houdt_van', 'Mens'),  ('Houdt_van', 'Student'), ('Kent', 'Mens')}, 
        'range':set() } # negatieve voorbeelden
       )
       ]
print(complete(RDFS))
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')}, 'subClassOf': {('Mens', 'Student'), ('Dier', 'Mens'), ('Student', 'Dier'), ('Mens', 'Dier'), ('Student', 'Mens'), ('Dier', 'Student')}, 'domain': {('Houdt_van', 'Mens'), ('Kent', 'Mens'), ('Houdt_van', 'Dier'), ('Houdt_van', 'Student')}, 'range': {'Vindt_lekker', 'Eten'}}
```

```python

```