# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Mustafa| secret |
|Thijs|secret |
|Jonne| secret |
| | |

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-1.1.1"><span class="toc-item-num">1.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-1.1.2"><span class="toc-item-num">1.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-1.1.2.1"><span class="toc-item-num">1.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-1.1.2.2"><span class="toc-item-num">1.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-1.1.2.3"><span class="toc-item-num">1.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-4.0.1"><span class="toc-item-num">4.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
studnr=12890111 # verander door het student id van jullie groepsleider
functienaam= 'complete_'+str(studnr)
print(functienaam)
# jouw functie


def complete_12890111(Facts, Rules): 
    Total_facts = Facts
    New_facts = Facts
    while New_facts:
        Round_facts = set()
        for triple in New_facts:
            for rule in Rules['subPropertyOf']:
                if triple[1] == rule[0]:
                    Round_facts.add((triple[0], rule[1], triple[2]))
            for rule in Rules['subClassOf']:
                if triple[2] == rule[0]:
                    Round_facts.add((triple[0], triple[1], rule[1]))
            for rule in Rules['domain']:
                if triple[1] == rule[0]:
                    Round_facts.add((triple[0], 'isa', rule[1]))
            for rule in Rules['range']:
                if triple[1] == rule[0]:
                    Round_facts.add((triple[2], 'isa', rule[1]))
        New_facts = Round_facts
        Total_facts = Total_facts | Round_facts
        
    return Total_facts




# test
# facts= test1[0][0]
# rules= test1[0][1]
# print(facts)
# print(rules)
# functienaam(facts,rules)
```

Output:
```
complete_12890111
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': {('Vindt_lekker','Eten')}
      }
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-13-6fc02bc1cbf3> in <module>
      5       }
      6 RDF= {('Henk','isa','Mens')}
----> 7 completion= functienaam(RDF,RDFS)
      8 assert isinstance(completion,set)
      9 for t in completion:

TypeError: 'str' object is not callable
```

# Jouw lijst met testen

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(studnr)

# testnaam= 'test_'+str(studnr)

# TEST 1
RDFS= {'subPropertyOf': set(),
      'subClassOf': {('Student','Mens'),('Mens','Dier')},
      'domain': set(),
      'range': set()}

RDF = {('Jan', 'isa', 'Student')}

positieve = {('Jan','isa','Mens'),('Jan','isa','Dier')}
negatieve = {('Jan','isa','arbeider')}

test1 = (RDF, RDFS, positieve, negatieve)

# TEST 2
RDFS= {'subPropertyOf': {('isin', 'has_capital')},
      'subClassOf': {('capital', 'city')},
      'domain': {('neighbours', 'country'), ('isin', 'city'), ('has_capital', 'capital'), ('has_minister', 'city')},
      'range': {('neighbours', 'country'), ('isin', 'country'), ('has_capital', 'capital'), ('has_minister', 'human')}}

RDF = {('Turkey', 'isa', 'Country'), ('Istanbul', 'isa', 'capital'), ('Greece', 'isa', 'country'), ('Turkey', 'neighbours', 'Greece'), ('Turkey', 'has_capital', 'Ankara'), ('Ankara', 'has_minister', 'Mansur_Yavas'), ('Ankara', 'has_population', '5504000')}

positieve = {('Ankara', 'isa', 'city'), ('Ankara', 'isin', 'Turkey')}
negatieve = {('Ankara', 'neighbours', 'country')}

test2 = (RDF, RDFS, positieve, negatieve)


# TEST 3
RDFS = {'subPropertyOf': {('jaagt_op', 'eet'), ('eet', 'vermoordt')},
        'subClassOf': {('leeuw', 'grote_kat'), ('grote_kat', 'katachtige'), ('katachtige', 'dier')},
        'domain': set(),
        'range': set()}

RDF = {('mufasa', 'isa', 'leeuw'), ('mufasa', 'jaagt_op', 'pumba')}
positieve = {('mufasa', 'vermoordt', 'pumba'), ('mufasa', 'isa', 'dier')}
negatieve = {('pumba', 'isa', 'grote_kat')}

test3 =  (RDF, RDFS, positieve, negatieve)

# TEST 4
RDFS= {'subPropertyOf': {('Rijdt_een', 'Houdt_van')},
      'subClassOf': {('Mercedes','Auto'),('Mens','Dier')},
      'domain':{('Rijdt_een', 'Bestuurder')},
      'range': {('Rijdt_een', 'Auto')}}

RDF = {('jan', 'Rijdt_een', 'Auto')}

positieve = {('jan', 'isa', 'Bestuurder'), ('jan', 'Houdt_van', 'Mercedes'), ('jan', 'Houdt_van', 'Auto'), ('Mercedes', 'isa', 'Auto'), ('jan', 'Rijdt_een', 'Auto')}
negatieve = {('jan', 'isa', 'Auto'), ('jan', 'Houdt_van', 'Bestuurder')}

test4 = (RDF, RDFS, positieve, negatieve)

# TEST 5
RDFS= {'subPropertyOf': {('Kijkt', 'Vindt_leuk')},
      'subClassOf': {('Mocro_Maffia', 'Videoland')},
      'domain': set(),
      'range': {('Kijkt', 'Serie')}}

RDF = {('jan', 'Kijkt', 'Mocro_Maffia')}

positieve = {('Mocro_Maffia', 'isa', 'Serie'), ('jan', 'Kijkt', 'Serie'), ('jan', 'Vindt_leuk', 'Serie'), ('jan', 'Vindt_leuk', 'Mocro_Maffia')}
negatieve = {('jan', 'Kijkt', 'Vindt_leuk')}

test5 = (RDF, RDFS, positieve, negatieve)

# TEST 6
RDFS= {'subPropertyOf': {('Leest', 'Kent'), ('Laat_uit', 'Loopt_met')},
      'subClassOf': {('Baasje','Dierenvriend')},
      'domain': {('Loopt_met','Baasje')},
      'range': {('Laat_uit','Dier')}}

RDF = {('Kees', 'Leest', 'Harry_Potter'),('Kees', 'Laat_uit', 'Takkie')}

positive = {('Kees', 'isa', 'Dierenvriend'),('Kees', 'Kent', 'Harry_Potter'),('Kees', 'Loopt_met', 'Takkie')}
negative = {('Kees', 'isa', 'Dier'),('Takkie','Loopt_met', 'Kees')}

test6 = (RDF, RDFS, positieve, negatieve)

#TEST 7

RDFS= {'subPropertyOf': {('Verkoopt', 'Verbouwt')},
      'subClassOf': {('Gewassen','Voedsel'),('Boer','Mens')},
      'domain': {('Verbouwt','Boer')},
      'range': {('Verbouwt','Gewassen')}}

RDF = {('Roderick', 'Verkoopt', 'Aardappelen')}
       
positive = {('Roderick', 'isa', 'Boer'),('Roderick', 'isa', 'Mens'),('Aardappel', 'isa', 'Voedsel')}
negative = {('Roderick','Verbouwt','Boer'), ('Roderick', 'isa', 'Gewassen')}

test7 = (RDF, RDFS, positieve, negatieve)



# Lijst met testen
test_11309598 = [*test1, test2, test3, test4, test5, test6, test7] 

testnaam=[test_11309598] # zet hier je tests in
```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.