# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
-----------------------|-------------|
|Tijn|secret     |
|Merijn| secret    |
|Cesar|secret     |


## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-Kennis-en-Data" data-toc-modified-id="Assignment-Kennis-en-Data-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Assignment Kennis en Data</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Peerreview" data-toc-modified-id="Peerreview-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Peerreview</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.4"><span class="toc-item-num">1.4&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li><li><span><a href="#Na-het-inleveren!" data-toc-modified-id="Na-het-inleveren!-1.5"><span class="toc-item-num">1.5&nbsp;&nbsp;</span>Na het inleveren!</a></span></li></ul></li><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-2.1"><span class="toc-item-num">2.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-2.1.1"><span class="toc-item-num">2.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-2.1.2"><span class="toc-item-num">2.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-2.1.2.1"><span class="toc-item-num">2.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-2.1.2.2"><span class="toc-item-num">2.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-2.1.2.3"><span class="toc-item-num">2.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
studnr=12495441 # verander door het student id van jullie groepsleider
functienaam= 'complete_'+str(12495441)
print(functienaam)
# jouw functie
def complete_12495441(facts, rules):
    def triples(facts, rules):
        output = []

        output.extend(tuple(((f[0], f[1], r[1]) for r in rules['subClassOf'] for f in facts if r[0] == f[2] )))
        output.extend(tuple(((f[0], r[1],f[2]) for r in rules['subPropertyOf'] for f in facts if r[0] == f[1] )))
        output.extend(tuple(((f[0], 'isa', r[1]) for r in rules['domain'] for f in facts if r[0] == f[1] )))
        output.extend(tuple(((f[2], 'isa', r[1]) for r in rules['range'] for f in facts if r[0] == f[1] )))
        return output
    
    output = triples(facts, rules)
    output2 = facts
    while output2:
        for triple in triples(output2, rules):  
            output.append(triple)
        output2 = triples(output2, rules)
    
    return (set(output))
```

Output:
```
complete_12495441
{('Henk', 'isa', 'Sterfelijk')}
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-6-f784c402e67c> in <module>
      4       'range': ('Vindt_lekker','Eten')}
      5 RDF= {('Henk','isa','Mens')}
----> 6 completion= functienaam(RDF,RDFS)
      7 assert isinstance(completion,set)
      8 for t in completion:

TypeError: 'str' object is not callable
```

```python
def check(triple, data):
    if triple in functienaam(data[0][0], data[0][1]):
        return 'Positief!'
    else:
        return 'Negatief!'

```

# Jouw lijst met testen

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(12495441)

# jouw test

test_12495441=[ 
    
      # test if function finds subClasses
      ({('Pietje','isa','Kind')},
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Kind','Mens'), ('Mens', 'Zoogdier')},
      'domain': set(),
      'range': set()},
       {('Pietje','isa','Zoogdier')}, # positieve voorbeelden
       {('Pietje', 'isa', 'Man')} # negatieve voorbeelden
       ),
       
    # test if function finds Subproperties
    ({('Tijn','isa','Mens'), ('Tijn', 'Vriend_van', 'Merijn'), ('Merijn', 'Getrouwd_met', 'Cesar')},  # RDF store
       {'subPropertyOf': {('Vriend_van', 'Kent'), ('Getrouwd_met', 'Vriend_van')},  # RDFS store
      'subClassOf': set(),
      'domain': set(),
      'range':  set()},
       {('Merijn','Kent','Cesar')}, # positieve voorbeelden
       {('Tijn', 'Vriend_van', 'Cesar')} # negatieve voorbeelden
       ),
       

    # test if function deduces all domains
    ({('Tijn', 'Houdt_van', 'Merijn'), ('Merijn', 'Vriend_van', 'Cesar'), ('Cesar', 'Vriend_van', 'Tijn'), 
      ('Cesar', 'Houdt_van', 'Tijn')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': {('Houdt_van', 'Mens'), ('Vriend_van', 'Mens'), ('Vriend_van', 'Sociaal')},
      'range':  set()},
       {('Cesar','isa','Sociaal')}, # positieve voorbeelden
       {('Tijn', 'isa', 'Sociaal')} # negatieve voorbeelden
       ),
       

      # test if function deduces all ranges
      ({('Pietje','Gaat_naar','school')},
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': set(),
      'domain': set(),
      'range': {('Gaat_naar', 'Geopend')}},
       {('School','isa','Geopend')}, # positieve voorbeelden
       {('School', 'isa', 'Huis')} # negatieve voorbeelden
       ),
       
    
    # test if function is able to loop through RDFS statements several times
    ({('Henk','isa','Man'), ('Karel', 'isa', 'Zoogdier')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Man','Persoon'), ('Persoon','Mens'), ('Mens','Zoogdier'), ('Zoogdier','Dier'), ('Dier','Wezen'), 
                     ('Wezen','Levend')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Levend'), ('Karel','isa','Wezen')}, # positieve voorbeelden
       {('Karel', 'isa', 'Mens')} # negatieve voorbeelden
       ),
    
    # Can the functionKan de code van nieuwe triples uit domain en range nieuwe triples maken met de subClassOf regels?
    ({('Cesar','Ouder_van','Jan'), ('Cesar','Houdt_van','Jan')},  # RDF store
       {'subPropertyOf': {},  # RDFS store
      'subClassOf': {('Persoon','Mens'),('Mens','Sterfelijk'), ('Mens','Dier'), ('Dier','Wezen')},
      'domain': {('Houdt_van','Persoon')},
      'range':  {('Ouder_van','Mens')}},
       {('Jan','isa','Sterfelijk')}, # positieve voorbeelden
       {('Jan', 'isa', 'Persoon')} # negatieve voorbeelden
       ),
    
    # test if function can make new triples with domain and range rules when new triples from subPropertyOf are added
    ({('Cesar','Is getrouwd met','Karel')},  # RDF store
       {'subPropertyOf': {('Is getrouwd met','Houdt_van'),('Houdt_van','Kent')},  # RDFS store
      'subClassOf': {('Mens','Sterfelijk'), ('Mens','Dier'), ('Dier','Wezen')},
      'domain': {('Houdt_van','Mens')},
      'range':  {('Kent','Kennis')}},
       {('Cesar','isa','Wezen')}, # positieve voorbeelden
       {('Kare', 'isa', 'Wezen')} # negatieve voorbeelden
       ),
       

    # test if function can handle all types of input; strings,integers and floats
    ({(12,'isa','Integer'), (12.5,'isa','Float')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Integer','Getal'),('Float','Getal')},
      'domain': set(),
      'range':  set()},
       {(12,'isa','Getal'), (12.5,'isa','Getal')}, # positieve voorbeelden
       {(13, 'isa', 'Getal')} # negatieve voorbeelden
       ),

    # Small test with all 4 types of RDFS statements
    ({('Henk','isa','Mens'), ('Jan','Houdt_van','Henk'), ('Henk', 'Vindt_lekker', 'Pizza'), ('Tijn', 'Ouder_van', 'Merijn'),
     ('Cesar', 'Houdt_van', 'Henk'), ('Frits', 'Vriend_van', 'Destiny')},  # RDF store
       {'subPropertyOf': {('Houdt_van','Kent')},  # RDFS store
      'subClassOf': {('Mens','Sterfelijk'), ('Mens','Dier'), ('Dier','Wezen')},
      'domain': {('Houdt_van','Mens'), ('Vriend_van', 'Mens')},
      'range':  {('Vindt_lekker','Eten'), ('Ouder_van','Mens'), ('Vriend_van','Persoon')}},
       {('Henk','isa','Sterfelijk'), ('Jan', 'Kent', 'Henk'), ('Destiny', 'isa', 'Persoon')}, # positieve voorbeelden
       {('Henk', 'Houdt_van', 'Ingrid')} # negatieve voorbeelden
       ),
    
    # Big test with all 4 types of RDFS statements
    ({('Pietje','Houdt_van','School'), ('Pietje', 'Leert', 'Engels'), ('Pino', 'Vliegt', 'Dagelijks'),('Rutte', 'Vertegenwoordigd', 'VVD'),('Pietje', 'isa', 'Man'), ('Nemo', 'isa', 'Vis'), ('Nelly', 'Rijdt', 'Auto'),('Laurens', 'isa', 'Man'),('Laurens', 'Print', 'Huiswerk'), ('Laurens', 'Allergisch_voor', 'Kaas')},
       {'subPropertyOf': {('Allergisch_voor','Eet_geen'), ('Leert', 'Volgt'), ('Houdt_van','Gaat_naar'), ('Vertegenwoordigd', 'Werkt_voor')},  # RDFS store
      'subClassOf': {('Mens','Zoogdier'),('Man','Mens'), ('Engels', 'Talen'),('Vis', 'Dier'), ('Vogel', 'Dier'), ('Bestuurder', 'Mens'), ('Politicus', 'Mens')},
      'domain': {('Houdt_van','Mens'), ('Leert', 'Leerling'),('Vertegenwoordigd', 'Politicus'), ('Rijdt','Bestuurder'), ('Vliegt', 'Vogel') },
      'range': {('Gaat_naar', 'Geopend'), ('Houdt_van', 'Leuk'), ('Leert', 'Vak') ,('Print', 'Papier'), ('Vertegenwoordigd', 'Partij'), ('Rijdt', 'Voertuig')}},
       { ('Nemo', 'isa', 'Dier'),('Rutte', 'isa', 'Zoogdier'),('Engels', 'isa', 'Vak')}, # positieve voorbeelden
       {('Engels', 'isa', 'partij'), ('Pietje', 'isa', 'Dier'), ('Pietje', 'isa', 'Leuk')} # negatieve voorbeelden
       )
    
       ] 
```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-8-500b0aa181b3> in <module>
----> 1 assert isinstance(testnaam,list)
      2 for test in testnaam:
      3     assert isinstance(test,tuple)
      4     assert len(test)==4
      5     assert isinstance(test[0],set)

AssertionError:
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.

```python
def complete(RDFS):
    subproperty = list(RDFS['subPropertyOf'])
    subclass = list(RDFS['subClassOf'])
    domain = list(RDFS['domain'])
    range = list(RDFS['range'])
    subproperty_oud, subclass_oud, domain_oud, range_oud = set(subproperty), set(subclass), set(domain), set(range)
    
    subproperty.extend((rule[0], rule2[1]) for rule in subproperty for rule2 in subproperty if rule[1] == rule2[0])           
    subclass.extend((rule[0], rule2[1]) for rule in subclass for rule2 in subclass if rule[1] == rule2[0])                
    domain.extend((rule[0], rule2[1]) for rule in domain for rule2 in subclass if rule[1] == rule2[0])                
    domain.extend((rule2[0], rule[1]) for rule in domain for rule2 in subproperty if rule2[1] == rule[0])                
    range.extend((rule[0], rule2[1]) for rule in range for rule2 in subclass if rule[1] == rule2[0])                
    range.extend((rule2[0], rule[1]) for rule in range for rule2 in subproperty if rule2[1] == rule[0])
    
    return {'subPropertyOf': set(subproperty) - subproperty_oud, 'subClassOf':set(subclass) - subclass_oud, 
            'Domain':set(domain) - domain_oud, 'Range':set(range) - range_oud}
    
```

Volgens ons zijn er 6 manieren waarop er nieuwe regels afgeleid kunnen worden uit andere regels. Nieuwe subclass regels maken kan doordat er een nieuwe wordt gevormd uit twee andere regels in subclass. Dit kan als x een subclass is van y, y een subclass is van z dan kun je hieruit de nieuwe regel x is een subclass van z halen. 
Nieuwe subproperty regels kunnen gevormd worden uit andere subroperty regels. Deze manier lijkt erg op de manier zoals bij subclass, namelijk als x een relatie is tussen twee waarden en y daardoor ook, y een relatie is en daardoor z ook dan is z ook een relatie door x. Dus je krijgt dan de regel (x,z). 
Er zijn twee manieren om nieuwe domeinregels af te leiden. Eén manier is door te kijken naar waar de class waar je toe behoort volgens de regel allemaal subclass van is. Dus als je doordat je van iemand houdt een persoon bent en volgens subclass ben je omdat je een persoon bent ook een dier, is een nieuwe regel in domein dat je omdat je van iemand houdt een dier bent. De andere manier is door naar subproperties te kijken. Als je bijvoorbeeld hebt in subproperty als je studeert dan leer je en in domain als je leert ben je student, dan kun je dus in domein de regel maken dat als je studeert je student bent. 
Voor nieuwe range regels zijn ook twee manieren. Eén manier is door te kijken waar de class in de regel weer subclass van is in de subclass regels. Dus als de regel is: als je ouder_van iets bent, ben je een persoon en in de subclass regels staat: persoon is subclass van mens, dan zou een nieuwe regel in range dus zijn: als je ouder bent van iets, dan ben je mens, oftwel ('ouder_van', 'mens'). De ander manier is als er in een subproperty een regel staat zoals als je moeder bent van iets ben je ook ouder van dat iets ('moeder van','ouder_van') en vervolgens de relatie 'moeder van' de regels van 'ouder van' overneemt. Dus als in range staat als je ouder bent van iets dan bent je persoon ('ouder_van', 'mens') betekent dat dat er een nieuwe regel is die zegt dan is het ook zo dat: als je moeder bent van iets ben je een mens ('moeder van','mens') 
Aangezien dit volgens ons de manieren zijn waarop nieuwe regels gevonden kunnen worden, zou de code volledig zijn als deze manieren gebruikt worden. Alle zes manieren van nieuwe regels maken voert onze code uit, dus daarom is onze code volledig.

### Onze Tests


```python
testset_12495441 = [
    
({'subPropertyOf': {('Opa_van','Grootouder_van'), ('Grootouder_van','Voorouder_van')},  
              'subClassOf': set(),
              'domain': set(),
              'range':  set()}),
           
({'subPropertyOf': set(),
      'subClassOf': {('Kaas','Eten'),('Eten','Voedsel')},
      'domain': set(),
      'range': set()}),

({'subPropertyOf': {('Lijsttrekker_van','Vertegenwoordigt_partij')},
      'subClassOf': set(),
      'domain': {('Vertegenwoordigt_partij', 'Politicus')} ,
      'range': set()}),

({'subPropertyOf': set(),
      'subClassOf': {('Man','Persoon')},
      'domain': {('Homo_met', 'Man')} ,
      'range': set()}),

({'subPropertyOf': {('Bestuurt', 'Rijdt')},
      'subClassOf': set(),
      'domain': set(),
      'range': {('Rijdt', 'Voertuig')}}),

({'subPropertyOf': set(),
      'subClassOf': {('Eten', 'Voeding')},
      'domain': set(),
      'range': {('Eet_graag', 'Eten')}}),

({'subPropertyOf': {('Moeder van','Ouder_van'), ('Studeert','Leert'), ('Houdt_van','Kent naam'), ('Kent naam','Kent'),
                              ('Kent','Heb gezien')},  
              'subClassOf': {('Persoon','Mens'), ('Mens','Zoogdier'), ('Zoogdier','Dier'), ('Dier','Wezen'), 
                             ('Wezen','Levend')},
              'domain': {('Houdt_van','Persoon'), ('Leert','Student')},
              'range':  {('Ouder_van', 'Persoon')}})
]
```