# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Joris| secret |
|Juriaan| secret|
|Mijntje| secret|
|Rik| secret|

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-Kennis-en-Data" data-toc-modified-id="Assignment-Kennis-en-Data-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Assignment Kennis en Data</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Peerreview" data-toc-modified-id="Peerreview-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Peerreview</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.4"><span class="toc-item-num">1.4&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li><li><span><a href="#Na-het-inleveren!" data-toc-modified-id="Na-het-inleveren!-1.5"><span class="toc-item-num">1.5&nbsp;&nbsp;</span>Na het inleveren!</a></span></li></ul></li><li><span><a href="#RDF-en-RDFS-completion" data-toc-modified-id="RDF-en-RDFS-completion-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>RDF en RDFS completion</a></span><ul class="toc-item"><li><span><a href="#Jouw-taak" data-toc-modified-id="Jouw-taak-2.1"><span class="toc-item-num">2.1&nbsp;&nbsp;</span>Jouw taak</a></span><ul class="toc-item"><li><span><a href="#Je-cijfer-is-gebaseerd-op" data-toc-modified-id="Je-cijfer-is-gebaseerd-op-2.1.1"><span class="toc-item-num">2.1.1&nbsp;&nbsp;</span>Je cijfer is gebaseerd op</a></span></li><li><span><a href="#Voorwaarden" data-toc-modified-id="Voorwaarden-2.1.2"><span class="toc-item-num">2.1.2&nbsp;&nbsp;</span>Voorwaarden</a></span><ul class="toc-item"><li><span><a href="#Database-en-regels" data-toc-modified-id="Database-en-regels-2.1.2.1"><span class="toc-item-num">2.1.2.1&nbsp;&nbsp;</span>Database en regels</a></span></li><li><span><a href="#Test-voorbeelden" data-toc-modified-id="Test-voorbeelden-2.1.2.2"><span class="toc-item-num">2.1.2.2&nbsp;&nbsp;</span>Test voorbeelden</a></span></li><li><span><a href="#Check-ajb" data-toc-modified-id="Check-ajb-2.1.2.3"><span class="toc-item-num">2.1.2.3&nbsp;&nbsp;</span>Check ajb</a></span></li></ul></li></ul></li></ul></li><li><span><a href="#Jouw-completion-functie" data-toc-modified-id="Jouw-completion-functie-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Jouw completion functie</a></span></li><li><span><a href="#Jouw-lijst-met-testen" data-toc-modified-id="Jouw-lijst-met-testen-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Jouw lijst met testen</a></span></li><li><span><a href="#Voor-de-echte-diehards-die-gaan-voor-een-10" data-toc-modified-id="Voor-de-echte-diehards-die-gaan-voor-een-10-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Voor de echte diehards die gaan voor een 10</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?" data-toc-modified-id="Hoeveel-RDFS-statements-zijn-er-eigenlijk-die-zouden-kunnen-volgen?-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?</a></span></li></ul></li></ul></li></ul></div>

# RDF en RDFS completion

* We maken een programma dat gegeven een RDF triple store plus een set RDFS statements, alle mogelijke niet triviale logische consequenties hiervan afleidt, en die in een RDF en RDFS database zet.
* Dit wordt ook wel *materialization* genoemd. Het idee is dat uit beide kennisbanken **exact dezelfde feiten volgen**, maar dat dat sneller kan vanuit de pure RDF versie (omdat je dan geen feiten hoeft af te leiden).
* De technische term voor het algorithme dat deze "completering" doet is de *chase*.
* Het [artikel van Enrico Franconi et al](https://link.springer.com/chapter/10.1007/978-3-642-41335-3_7) bevat een systeem.
* Net als in Franconi werken we alleen met de 4 interresantste RDFS statements:
    * **Hierarchie** 
        * $\forall xy (Rxy \rightarrow Sxy)$  (" to love him is to know him", als je van iemand houdt dan ken je diegene) 
        * $\forall x (Horse(x) \rightarrow Animal(x))$
    * **Domain en range restrictions**
        * $\forall xy (Rxy \rightarrow Dx)$ (domain) "Als x een staart heeft is het een dier"
        * $\forall xy (Rxy \rightarrow Dy)$ (range)  "Als je een  ouder bent van y, is y een mens."
        
## Jouw taak

1. Maak de functie `complete(facts, rules)` die voldoet aan de volgende eigenschap:
    * voor elke feit `(o,p,s)`, facts+rules $\models$ (o,p,s) $\Leftrightarrow$ `complete(facts, rules)` $\models$ (o,p,s) 
        * Jouw functie is zowel volledig als geldig (*sound and complete*), dat is exact wat je wilt.
    * `complete(facts, rules)` is een verzameling feiten.
2. Maak een flinke collectie van tests waarmee je test dat jouw functie aan die eisen voldoet. 
    * Dus dat het niet te veel en niet te weinig, maar exact de juiste feiten impliceert.
    * Je deelt die tests met alle anderen, en ook jouw functie, en we gaan dus elke functie tegen elke test draaien, en dus een matrix met scores maken.
    
### Je cijfer is gebaseerd op 

1. Hoe goed doorstaat jouw functie alle tests.
2. Hoe goed zijn jouw tests.
    * Strenge (doch rechtvaardige) tests zijn natuurlijk het best. Tests waar iedereen voor slaagt lijken leuk maar zijn meestal niet veel waard. 
    
    
### Voorwaarden

#### Database en regels

We werken niet direct met RDF en RDFS omdat dat zo'n enorme syntax overhead heeft. We doen de dingen gewoon lekker in Python.

* Je RDF store is een verzameling triples, waarbij de 3 elementen strings of integers of floats zijn.
    * Je schrijft eigenschappen van individuen als `('henk', 'isa', 'Man')`
        * dus altijd met de vaste relatie `isa` in het midden.
    * en relaties tussen individuen als `('henk', Houdt_van', 'ingrid')`
    * en idem voor relaties tussen individuen en getallen `('henk', 'leeftijd', 53)`
    * Begin namen van individuen ("knopen") meteen kleine letter, en eigenschappen en relaties met een grote.
* Je RDFS store is een dict met sleutels `['subPropertyOf', 'subClassOf', 'domain', 'range']` en voor elke sleutel een verzameling met 2-tuples (paartjes in gewoon Nederlands) als waarde. 

#### Test voorbeelden

Een testset is een lijst triples van de vorm `(D:RDF_store, R:RDFS_store, Pos:RDF_store, Neg:RDF_store)` waarvoor geldt dat 
* voor elk triple t in Pos: $D,R\models t$
* voor elk triple t in Neg: $D,R\not\models t$
* Met $\models$ wordt dus bedoeld volgens de Franconi semantiek, en dus ook volgens de RDF semantiek. Dat betekent dus niet per se dat dit hetzelfde is als volgens jouw completion functie (want die zou fouten kunnen maken).
* Hieronder staat een flauw voorbeeld. Jounes is veel gaver en strenger en leuker  natuurlijk!

#### Check ajb

* Als de tests hieronder niet slagen kunnen we niks met je code, en kan je  dus niet meedoen. Zonde van al je werk toch?

```python
# voorbeeld
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDFS
```

Output:
```
{'subPropertyOf': {('Eet', 'Vindt_lekker'), ('Houdt_van', 'Kent')},
 'subClassOf': {('Mens', 'Dier'), ('Student', 'Mens')},
 'domain': {('Houdt_van', 'Mens')},
 'range': ('Vindt_lekker', 'Eten')}
```

```python
# voorbeeld

test1= [
    ({('Henk','isa','Mens')},  # RDF store
       {'subPropertyOf': set(),  # RDFS store
      'subClassOf': {('Mens','Sterfelijk')},
      'domain': set(),
      'range':  set()},
       {('Henk','isa','Sterfelijk')}, # positieve voorbeelden
       {('henk', 'Houdt_van', 'ingrid')} # negatieve voorbeelden
       )
       ]
test1
```

Output:
```
[({('Henk', 'isa', 'Mens')},
  {'subPropertyOf': set(),
   'subClassOf': {('Mens', 'Sterfelijk')},
   'domain': set(),
   'range': set()},
  {('Henk', 'isa', 'Sterfelijk')},
  {('henk', 'Houdt_van', 'ingrid')})]
```

# Jouw completion functie

```python
# a color class used to clarify output
class color:
    CYAN = '\033[96m'
    GREEN = '\033[92m'
    BOLD = '\033[1m'
    END = '\033[0m'
```

<!-- grade_id: f, points: 0, grade: false, solution: true -->
```python
studnr=11309946 # verander door het student id van jullie groepsleider
functienaam= 'complete_'+str(studnr)
print(functienaam)

def subClassOf(RDF, RDFS):
    # (A,sc,B) & (X,type,A) == (X,type,B)
    return {(triple[0], 'isa', sc[1]) for sc in RDFS['subClassOf'] for triple in RDF if triple[2] == sc[0]}
    
def subPropertyOf(RDF, RDFS):
    # (A,sp,B) & (X,A,Y) == (X,B,Y)
    return {(triple[0], sp[1], triple[2]) for sp in RDFS['subPropertyOf'] for triple in RDF if triple[1] == sp[0]}

def Domain(RDF, RDFS):
    # (A,dom,B) & (X,A,Y) == (X,type,B)
    return {(triple[0], 'isa', dom[1]) for dom in RDFS['domain'] for triple in RDF if triple[1] == dom[0]}

def Range(RDF, RDFS):
    # (A,range,B) & (X,A,Y) == (Y,type,B)
    return {(triple[2], 'isa', ran[1]) for ran in RDFS['range'] for triple in RDF if triple[1] == ran[0]}
    
def complete_11309946(facts, rules):  
    completion, previous = facts, set()
    while completion != previous:
        previous = completion
        completion = completion | subClassOf(completion, rules) | subPropertyOf(completion, rules) | Domain(completion, rules) | Range(completion, rules)
    return completion

#test
complete_11309946(test1[0][0],test1[0][1])
```

Output:
```
complete_11309946
{('Henk', 'isa', 'Mens'), ('Henk', 'isa', 'Sterfelijk')}
```

<!-- grade_id: s, points: 0, grade: true, solution: false -->
```python
assert studnr!= 1234567, 'Dit is niet jouw student nummer!'

```

<!-- grade_id: ft, points: 0, grade: true, solution: false -->
```python
RDFS= {'subPropertyOf': {('Houdt_van','Kent'), ('Eet','Vindt_lekker')},
      'subClassOf': {('Mens','Dier'),('Student','Mens')},
      'domain': {('Houdt_van','Mens')},
      'range': ('Vindt_lekker','Eten')}
RDF= {('Henk','isa','Mens')}
completion= functienaam(RDF,RDFS)
assert isinstance(completion,set)
for t in completion:
    assert isinstance(t,tuple)
    assert len(t)==3
    assert isinstance(t[0], str)
    assert isinstance(t[1], str)
    assert isinstance(t[2], (str, int,float))
    
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-51-f784c402e67c> in <module>
      4       'range': ('Vindt_lekker','Eten')}
      5 RDF= {('Henk','isa','Mens')}
----> 6 completion= functienaam(RDF,RDFS)
      7 assert isinstance(completion,set)
      8 for t in completion:

TypeError: 'str' object is not callable
```

# Jouw lijst met testen

<!-- grade_id: t, points: 0, grade: false, solution: true -->
```python
testnaam= 'test_'+str(studnr)

# jouw test
test_11309946= [
    
        # TEST 1 - Empty RDF Store
        ({('henk','isa','Student')},  # RDF store
        {'subPropertyOf': set(),  # RDFS store
        'subClassOf': set(),
        'domain': set(),
        'range':  set()}, # RDFS store
        {('henk', 'isa', 'Student')}, # positief
        {('henk', 'isa', 'Mens')} # negatief
        ),
    
        # TEST 2
        ({('henk','isa','Student')},  # RDF store
        {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Mens','Sterfelijk'), ('Student','Mens')},
        'domain': set(),
        'range':  set()},
        {('henk', 'isa', 'Student'), ('henk', 'isa', 'Mens'), ('henk', 'isa', 'Sterfelijk')}, # positief
        {('henk', 'isa', 'Wezen')} # negatief
        ),
    
        # TEST 3
        ({('henk','isa','Mens')},  # RDF store
        {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Mens','Sterfelijk')},
        'domain': set(),
        'range':  set()},
        {('henk', 'isa', 'Mens'), ('henk', 'isa', 'Sterfelijk')}, # positief
        {('henk', 'Houdt_van', 'ingrid')} # negatief
        ),
    
        # TEST 4
        ({('henk', 'HeeftLesVan', 'marx')},  # RDF store
        {'subPropertyOf': set(),  # RDFS store
        'subClassOf': set(),
        'domain': {('HeeftLesVan', 'Student')},
        'range':  {('HeeftLesVan', 'Docent')}},
        {('henk', 'isa', 'Student'), ('marx', 'isa', 'Docent'), ('henk', 'HeeftLesVan', 'marx')}, # positief
        set() # negatief
        ),

        # TEST 5
        ({('henk','Leeftijd',18)},  # RDF store
        {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Mens','Sterfelijk')},
        'domain': {('Leeftijd','Mens')},
        'range':  set()},
        {('henk', 'Leeftijd', 18), ('henk', 'isa', 'Mens'), ('henk', 'isa', 'Sterfelijk')}, # positief
        {('henk','Leeftijd',50)} # negatief
        ),

        # TEST 6
        ({('henk','Leeftijd',18)},  # RDF store
        {'subPropertyOf': set(),  # RDFS store
        'subClassOf': {('Mens','Wezen'), ('Wezen','Sterfelijk')},
        'domain': {('Leeftijd','Mens')},
        'range':  set()},
        {('henk','isa','Sterfelijk'), ('henk','isa','Mens'), ('henk','isa','Wezen'), ('henk','Leeftijd',18)}, # positief
        {('henk','Leeftijd',50)} # negatief
        ),

        # TEST 7
        ({('uva','isa','universiteit'), ('marx','GeeftVak','knowledge_and_data')},  # RDF store
        {'subPropertyOf': {('GeeftVak','HeeftKennisVan')},  # RDFS store
        'subClassOf': {('docent', 'mens'), ('universiteit','onderwijsinstelling'), ('hogeschool','onderwijsinstelling'), ('docent','personeelslid')},
        'domain': {('GeeftVak', 'docent')},
        'range':  {('GeeftVak','vak')}},
        {('knowledge_and_data', 'isa', 'vak'), ('marx', 'isa', 'mens'), ('marx', 'isa', 'personeelslid'), 
         ('marx', 'HeeftKennisVan', 'knowledge_and_data'), ('uva', 'isa', 'universiteit'), 
         ('uva', 'isa', 'onderwijsinstelling'), ('marx', 'isa', 'docent'), ('marx', 'GeeftVak', 'knowledge_and_data')}, # positief
        {('uva','isa','hogeschool'), ('knowledge_and_data','isa','docent'), ('marx','isa','vak')} # negatief
        )
]

for i, test in enumerate(test_11309946, 1):
    print('\n---' + color.BOLD + 'Test ' + str(i) + color.END + '---')

    output = complete_11309946(test[0], test[1])
    print(f"Pos: {output == test[2]}") # True
    print(f"Neg: {output == test[3]}") # False
    print(f"Output: {output}")

```

Output:
```
---[1mTest 1[0m---
Pos: True
Neg: False
Output: {('henk', 'isa', 'Student')}

---[1mTest 2[0m---
Pos: True
Neg: False
Output: {('henk', 'isa', 'Mens'), ('henk', 'isa', 'Student'), ('henk', 'isa', 'Sterfelijk')}

---[1mTest 3[0m---
Pos: True
Neg: False
Output: {('henk', 'isa', 'Mens'), ('henk', 'isa', 'Sterfelijk')}

---[1mTest 4[0m---
Pos: True
Neg: False
Output: {('henk', 'isa', 'Student'), ('marx', 'isa', 'Docent'), ('henk', 'HeeftLesVan', 'marx')}

---[1mTest 5[0m---
Pos: True
Neg: False
Output: {('henk', 'isa', 'Mens'), ('henk', 'isa', 'Sterfelijk'), ('henk', 'Leeftijd', 18)}

---[1mTest 6[0m---
Pos: True
Neg: False
Output: {('henk', 'isa', 'Mens'), ('henk', 'isa', 'Sterfelijk'), ('henk', 'isa', 'Wezen'), ('henk', 'Leeftijd', 18)}

---[1mTest 7[0m---
Pos: True
Neg: False
Output: {('marx', 'GeeftVak', 'knowledge_and_data'), ('marx', 'isa', 'personeelslid'), ('marx', 'isa', 'docent'), ('marx', 'isa', 'mens'), ('uva', 'isa', 'universiteit'), ('marx', 'HeeftKennisVan', 'k
```

<!-- grade_id: tt, points: 0, grade: true, solution: false -->
```python
assert isinstance(testnaam,list)
for test in testnaam:
    assert isinstance(test,tuple)
    assert len(test)==4
    assert isinstance(test[0],set)
    assert isinstance(test[1],dict)
    assert isinstance(test[2],set)
    assert isinstance(test[3],set)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-53-500b0aa181b3> in <module>
----> 1 assert isinstance(testnaam,list)
      2 for test in testnaam:
      3     assert isinstance(test,tuple)
      4     assert len(test)==4
      5     assert isinstance(test[0],set)

AssertionError:
```

# Voor de echte diehards die gaan voor een 10

Franconi beschrijft in zijn artikel dat je ook RDFS statements uit andere RDFS statements kunt afleiden. Hij geeft een leuk voorbeeld in het eerste plaatje.

Dat kan jij natuurlijk ook!

Maak een functie `complete(RDFS)` die precies alle RDFS statements afleidt die volgen uit de verzameling `RDFS`.

Maak er ook weer een mooie testset bij, en beargumenteer dat jouw systeem geldig en volledig is.

### Hoeveel RDFS statements zijn er eigenlijk die zouden kunnen volgen?

* Gegeven de relaties en eigenschappen in de verzameling `RDFS` is dit eindig.
* En er zijn maar 4 soorten RDFS statements. 
* Dus moet het wel eindig zijn, en kan je zelfs uitrekenen hoeveel.
* Een manier is om ze allemaal te proberen....
    * Voeg het antecedent toe als RDF-triple aan `RDFS` en kijk of het consequent dan volgt...
* Reken het eens uit voor $r$ relaties en $p$ eigenschappen.
    * Met $r$ relaties kan je $r\cdot(r-1)$ veel `subPropertyOf` relaties maken toch...
    * en dan nog ....
    * en ....
* Dit is niet eens zo gek, maar met een *chase* gaat het natuurlijk veel sneller omdat je dan resultaatgericht werkt in plaats van als een dolle alles gaat proberen.
    * Bijv; Stel je hebt 100 relaties, maar in je RDFS komt er maar 1 voor in antecdent positie in `subPropertyOf` statements. Hoeveel `subPropertyOf` kan je dan afleiden? Ja, dat is dus sneller dan al die 100*99 te gaan proberen toch? Want dat is best veel.

```python
def AllsubPropertyOf(RDFS):
    return RDFS['subPropertyOf'] | {(sp[0], sp2[1]) 
                                    for sp in RDFS['subPropertyOf'] 
                                    for sp2 in RDFS['subPropertyOf'] 
                                    if sp[1] == sp2[0]}
def AllsubClassesOf(RDFS):
    return RDFS['subClassOf'] | {(sp[0], sp2[1]) 
                                 for sp in RDFS['subClassOf'] 
                                 for sp2 in RDFS['subClassOf'] 
                                 if sp[1] == sp2[0]}

def AllDomains(RDFS):
    return RDFS['domain'] | {(dom[0], dom2[1])
                                if dom[1] == dom2[0]
                                else (dom2[1], dom[1])
                                for dom in RDFS['domain']
                                for dom2 in (RDFS['subClassOf'] | RDFS['subPropertyOf'])
                                if dom[0] == dom2[0]}

def AllRanges(RDFS):
    return RDFS['range'] | {(ran[0], ran2[1])
                               if ran[1] == ran2[0]
                               else (ran2[1], ran[1])
                               for ran in RDFS['range']
                               for ran2 in (RDFS['subClassOf'] | RDFS['subPropertyOf'])
                               if ran[0] == ran2[0]}

def complete(RDFS):  
    completion, previous = RDFS, dict()
    while completion != previous:
        previous = completion
        completion = {'subPropertyOf': AllsubPropertyOf(completion), 
                     'subClassOf': AllsubClassesOf(completion),
                     'domain': AllDomains(completion), 
                     'range': AllRanges(completion)}
    return {k:completion[k] - RDFS[k] for k,v in completion.items()}     

```

## Toelichting op tests

De functie `complete(RDFS)` geeft uit een als parameter meegegeven dict van RDFS relaties terug welke andere relaties er mogelijkerwijs uit deze relaties zouden kunnen volgen. De aan de functie meegegeven RDFS-dict heeft als keys de 4 verschillende typen RDFS relaties zoals beschreven door Franconi (subproperties, subclasses, domains en ranges), en als values de in de RDFS bestaande relaties die onder dat type vallen. 

In onderstaand codeblok wordt de toepassing van de functie getest met een testset die een politiek thema heeft. Hieruit blijkt dat de functie alle mogelijk geïmpliceerde RDFS relaties returnt. Net als de meegegeven RDFS wordt de output van de functie gereturned in de vorm van een dict met de 4 verschillende typen RDFS relaties als keys. Hieronder een toelichting van het uitgevoerde voorbeeld op de testset.  

`Lijsttrekker` is een subproperty van `Kamerlid` en `Kamerlid` is een subproperty van `Vertegenwoordigt`. In ‘gewone mensentaal’ kan dit worden uitgesproken als “iedereen die een lijsttrekker is, is een kamerlid, en iedereen die een kamerlid is, vertegenwoordigt”. Aangezien `Lijsttrekker` een subproperty is van een subproperty van `Vertegenwoordigt`, is `Lijsttrekker` ook een subproperty van `Vertegenwoordigt`. Dit volgt ook uit de output. Voor subclasses geldt hetzelfde; `min_pres` is een subklasse van een subklasse van `kabinetslid`, waaruit volgt dat `min_pres` zelf ook een subklasse van `kabinetslid` is. 
 
Bij `domain` stelt de relatie `('Kamerlid', 'politicus')` dat alleen politici kamerleden kunnen zijn. Aangezien `politicus` een subklasse is van mens, volgt de domeinrelatie `(‘Kamerlid’, ‘mens’)` hieruit; hiermee wordt gesteld dat alleen mensen kamerleden kunnen zijn. `Kamerlid` is een subproperty van ‘Vertegenwoordigt’, derhalve volgt dat iedereen die `Vertegenwoordigt` een politicus moet zijn `(‘Vertegenwoordigt’, ‘politicus’)`. En, wederom vanuit de subklasse relatie `(‘politicus’, ‘mens’)` volgt hieruit ook `(‘Vertegenwoordigt’, ‘mens’)` als domeinrelatie. Tenslotte stelt de in de RDFS meegegeven range relatie `('Kamerlid', 'partij')` dat als iemand kamerlid is, die persoon dan kamerlid van een partij is. Aangezien `Kamerlid` een subproperty is van `Vertegenwoordigt` volgt hier ook uit dat hetgeen wat iemand vertegenwoordigt een partij is. 


```python
# Input
RDFS = {'subPropertyOf': {('Lijsttrekker', 'Kamerlid'), ('Kamerlid', 'Vertegenwoordigt')},
      'subClassOf': {('min_pres','minister'), ('minister', 'kabinetslid'), ('politicus', 'mens')},
      'domain': {('Kamerlid', 'politicus')},
      'range':  {('Kamerlid', 'partij')}}

# Output
print(color.BOLD + 'Test-output' + color.END + '\n')
for statement in RDFS:
    print(color.BOLD + statement + color.END)
    print(color.CYAN + 'Input: ' + str(RDFS[statement]) + color.END)
    print(color.GREEN + 'Output: ' + str(complete(RDFS)[statement]) + color.END)
    print()
```

Output:
```
[1mTest-output[0m

[1msubPropertyOf[0m
[96mInput: {('Lijsttrekker', 'Kamerlid'), ('Kamerlid', 'Vertegenwoordigt')}[0m
[92mOutput: {('Lijsttrekker', 'Vertegenwoordigt')}[0m

[1msubClassOf[0m
[96mInput: {('minister', 'kabinetslid'), ('min_pres', 'minister'), ('politicus', 'mens')}[0m
[92mOutput: {('min_pres', 'kabinetslid')}[0m

[1mdomain[0m
[96mInput: {('Kamerlid', 'politicus')}[0m
[92mOutput: {('Vertegenwoordigt', 'politicus')}[0m

[1mrange[0m
[96mInput: {('Kamerlid', 'partij')}[0m
[92mOutput: {('Vertegenwoordigt', 'partij')}[0m
```