# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
|Jesse| secret |
|Alex| secret |
|Jesse| secret |
|Thomas| secret |

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Q1-logisch-gevolg-(4pts)" data-toc-modified-id="Q1-logisch-gevolg-(4pts)-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Q1 logisch gevolg (4pts)</a></span></li><li><span><a href="#Q2-Vreemde-modellen-(4pts)" data-toc-modified-id="Q2-Vreemde-modellen-(4pts)-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Q2 Vreemde modellen (4pts)</a></span></li><li><span><a href="#Q3-Conjunctieve-vragen-aan-een-netwerk-(4pts)" data-toc-modified-id="Q3-Conjunctieve-vragen-aan-een-netwerk-(4pts)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Q3 Conjunctieve vragen aan een netwerk (4pts)</a></span></li><li><span><a href="#Q1-logisch-gevolg" data-toc-modified-id="Q1-logisch-gevolg-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Q1 logisch gevolg</a></span></li><li><span><a href="#Q2-Vreemde-modellen" data-toc-modified-id="Q2-Vreemde-modellen-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Q2 Vreemde modellen</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Deel-1-(2pt)" data-toc-modified-id="Deel-1-(2pt)-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Deel 1 (2pt)</a></span></li><li><span><a href="#Deel-2-(2pt)" data-toc-modified-id="Deel-2-(2pt)-5.0.2"><span class="toc-item-num">5.0.2&nbsp;&nbsp;</span>Deel 2 (2pt)</a></span></li><li><span><a href="#deel-1" data-toc-modified-id="deel-1-5.0.3"><span class="toc-item-num">5.0.3&nbsp;&nbsp;</span>deel 1</a></span></li><li><span><a href="#deel-2" data-toc-modified-id="deel-2-5.0.4"><span class="toc-item-num">5.0.4&nbsp;&nbsp;</span>deel 2</a></span></li></ul></li></ul></li><li><span><a href="#Q3-Conjunctieve-vragen-aan-een-netwerk-(4pt)" data-toc-modified-id="Q3-Conjunctieve-vragen-aan-een-netwerk-(4pt)-6"><span class="toc-item-num">6&nbsp;&nbsp;</span>Q3 Conjunctieve vragen aan een netwerk (4pt)</a></span><ul class="toc-item"><li><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Eisen" data-toc-modified-id="Eisen-6.0.0.1"><span class="toc-item-num">6.0.0.1&nbsp;&nbsp;</span>Eisen</a></span></li></ul></li></ul></li><li><span><a href="#Deel-2" data-toc-modified-id="Deel-2-6.1"><span class="toc-item-num">6.1&nbsp;&nbsp;</span>Deel 2</a></span></li></ul></li></ul></div>

```python
from typing import Generator
from nose.tools import assert_equal, assert_count_equal
import networkx as nx
%matplotlib inline
```

# Q1 logisch gevolg (4pts)

* Je maakt een Python programma waaraan je een verzameling premissen $\Sigma$ en een conclusie $\phi$ in propositielogische formules geeft en die teruggeeft of $\Sigma\models\phi$. 
    * In het geval van niet geef je een (of alle) tegenvoorbeelden terug.
    
# Q2 Vreemde modellen (4pts)

* Je maakt vreemde modellen voor je netwerk van familierelaties. Modellen die wel mogen, maar in het echt niet of vrijwel nooit voorkomen. 
* Daarna probeer je die met logische formules te verbieden. Je verzint dus zinnen die *niet waar* zijn op die ongewenste modellen.

# Q3 Conjunctieve vragen aan een netwerk (4pts)

* Een netwerk met feiten (en regels) is natuurlijk ook een database (of een *triple store*), en wat doen we met dat soort dingen? We stellen er vragen aan, en krijgen dan antwoord. Normaal doe je dat in SQL. In deze cursus leer je hoe je dat doet in SPARQL. Later bij Data Science leer je hoe je dat doet met `pandas`. 
* De preciese taal maakt niet, de meeste vragen die je stelt in je leven (en alle andere mensen op aarde) zijn zogenaamde *conjunctive queries*: eerste orde formules van de vorm $\exists \bar{x} (\bigwedge R(x_i,x_j))$. 
    * OK OK, je gebruikt ook nog wel eens het connectief "of", maar dat is makkelijk toe te voegen aanje programma als je dit hebt.
        * (want je kunt alle "of"s naar buiten brengen, en dan heb je dus een grote disjunctie van CQs.
        * die draai je allemaal achter (of naast) elkaar, en neemt de vereniging van de antwoorden.
        * Probeer maar eens een voorbeeld
* Bijvoorbeeld, "wie zijn de grootouders van Henkie?" wordt dan  $\exists z (R(x,z) \wedge R(z,"Henkie"))$. Je kunt dat zelf vast prima in SQL en volgende week in SPARQL zeggen.

# Q1 logisch gevolg

* We werken met de [sympy module](https://docs.sympy.org/latest/modules/logic.html). Dat is erg handig. Neem hem eerst goed door en speel er wat mee.
* In de cel hieronder vind je een handige routine.
* Je gebruikt natuurlijk  de definitie van $\Sigma\models\phi$, van logisch gevolg dus.
* En daarvoor is de `sympy` routine `satisfiable` best handig....
* De tests in de antwoord cel gaven bij mij dit:
<pre>
{p, q | ~p} q True
{q, q | ~p} p [{q: True, p: False}]
{q, Implies(p, q)} p {q: True, p: False}
{q & (r | s), Implies(p, q)} p [{q: True, r: True, s: True, p: False}, {q: True, r: True, s: False, p: False}, {q: True, s: True, r: False, p: False}]
WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.
</pre>

```python
from sympy import *
from sympy.logic.inference import satisfiable

p,q,r,s,t = symbols('p q r s t')

def maakconjunctie(P):
    '''Creates a conjunction from a set or list of formulas.'''
    P=list(P)
    if P:
        return P[0]&maakconjunctie(P[1:])
    else:
        return True
#test
maakconjunctie({p, ~p |q, p|(p&q)})    
```

Output:
```
p & (q | ~p) & (p | (p & q))
```

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
import random
import itertools
from sympy.logic.boolalg import Implies

def implies(Premisses, Conclusion, allmodels=False):
    '''Return True if Conclusion follows from Premisses,
    return a (or all) countermodel(s) otherwise. '''
    P= maakconjunctie(Premisses)

    if not Conclusion.atoms().issubset(P.atoms()):
        print('WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.')
    else:
        S=list(satisfiable(P,all_models=True))
        C = list(satisfiable(Conclusion,all_models=True))

        if S[0] and C[0]:
            testlist = []
            for y in C:
                for z in y:
                    for x in S:
                        testlist.append(x[z] == y[z])
            if all(testlist):
                return True

            # if all([x[z] == y[z] for x in S for z in y for y in C]):
                # return True
            else:
                modellist = []
                for y in C:
                    for z in y:
                        for x in S:
                            if x[z] != y[z]:
                                modellist.append(x)
                if allmodels:
                    return modellist
                    # return [x for x in S for z in y for y in C if x[z] != y[z]]
                else:
                    return modellist[0]
                    # return [x for x in S for z in y for y in C if x[z] != y[z]][0]
        else:
            print("WARNING: premisses and conclusion must be satisfiable")



# TEST
P= {~p|q,p}
C= q

print(P, C, implies(P,C, allmodels=True ))
P= {~p|q,q}
C= p

print(P, C, implies(P,C, allmodels=True ))

P= {p>>q,q}  # andere manier om implicatie te schrijven
C= p

print(P, C, implies(P,C))

P= {p>>q,q,(r|s)}  # andere manier om implicatie te schrijven
C= p

print(P, C, implies(P,C, allmodels=True ))
P= {p>>q,q}  # andere manier om implicatie te schrijven
C= r

implies(P,C) # Dit geeft een error. Probeer het ook zelf
```

Output:
```
{p, q | ~p} q True
{q, q | ~p} p [{q: True, p: False}]
{q, Implies(p, q)} p {q: True, p: False}
{r | s, q, Implies(p, q)} p [{q: True, s: True, r: True, p: False}, {q: True, s: True, r: False, p: False}, {q: True, r: True, s: False, p: False}]
WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.
```

<!-- grade_id: q1t, points: 4, grade: true, solution: false -->
```python
assert_equal(type(implies({p},p)), bool)
assert_equal(type(implies({p|q},p)), dict)
assert_equal(type(implies({p|q},p, allmodels=True)), list)

```

# Q2 Vreemde modellen

### Deel 1 (2pt)
1. Je maakt vreemde modellen voor je netwerk van familierelaties. Modellen die wel mogen, maar in het echt niet of vrijwel nooit voorkomen. 
2. Daarna probeer je die met logische formules te verbieden. Je verzint dus zinnen die *niet waar* zijn op die ongewenste modellen.

In ieder geval verzin je zinnen die de volgende vijf modellen verbieden (die soort $\Sigma$ bij een knoop is eigenlijk een M en staat voor Man. Gebruik dus weer de predicaatlogische relaties $R,M,V$ in die zinnen, en niks anders. Geef duidelijk aan welk model je met welke zin verbiedt.

![](vreemd.jpg)

### Deel 2 (2pt)

Ga nu verder en verzin nog meer modellen die je zou willen verbieden. Misschien vindt je er wel oneindig veel! Kan je daar een zin die die allemaal verbiedt voor vinden? Kan je het met oneindig veel zinnen? Geef ze!

**Je antwoorden zijn keurige Latex formules, en natuurlijk wel gevormd, met nette haakjes etc.**

**Voor je modellen, maak je plaatjes in netwerkx, met de hand, of iets anders, als het maar heel duidelijk is (en dus niet zo als die in deze opgave ;-)**

<!-- grade_id: q2a, points: 4, grade: true, solution: true -->
### Definitie primitieven

M = Man

V = Vrouw

R = Ouder-kindrelatie



### Antwoord deel 1
Volgorde onderstaande twee verzamelingen regels slaan op dezelfde regels in beide verzamelingen
De locatie van de desbetreffende regel in de bovenstaande afbeelding is weergegeven na de regel tussen haakjes.
###### Regels uit het plaatje:
1. Iemand heeft twee ouders hebben die allebei man zijn (linksboven)
2. Iemand is ouder van zichzelf (rechtsboven)
3. Iemand kan niet zowel de opa/oma als de vader/moeder zijn van iemand (midden)
4. Iemand is vader/moeder van zijn of haar opa/oma (linksonder)
5. Iemand is vader/moeder van zijn of haar vader/moeder (rechtsonder)

###### Regels die deze zinnen uit het plaatje verbieden
 | Formule | Uitleg |
 | ------- | ------ |
 | $\neg \exists xyz((R(y,x) \land M(y)) \land (R(z,x) \land M(z)) \land \neg(y=z)+ )$ | Niemand heeft twee ouders die beiden man zijn (linksboven) |
 | $\neg \exists xy(R(x,y) \land (x=y))$ | Niemand is ouder van zichzelf (rechtsboven) |
 | $\neg \exists xyz(R(x,z) \land R(z,y) \land R(x,y))$ | Iemand kan niet zowel de opa/oma als de vader/moeder zijn van iemand (midden) |
 | $\neg \exists xyz(R(x,z) \land R(z,y) \land R(y,x))$ | Niemand is ouder van zijn of haar oma/opa (linksonder) |
 | $\neg \exists xy(R(x,y) \land R(y,x))$  | Niemand is ouder van zijn of haar vader/moeder (rechtsonder) |
 
 
### Antwoord deel 2
Zin die voor alles verbiedt:
 | Formule | Uitleg |
 | ------- | ------ |
 | $\forall xyz( Kind (x,y) \leftrightarrow (R(y,x) \wedge \neg (x = y)) \wedge (R(z,x) \wedge \lnot (x = z)) \wedge (M(z) \wedge V(y)) \veebar (M(y) \wedge V(z)) )$ | Je kan alleen kind van je vader en moeder zijn |


###### Regels buiten het plaatje die vreemd zijn
1. Iemand is ouder van zijn of haar broer/zus
2. Iemand is ouder van zijn of haar overgrootvader/overgrootmoeder
3. Iemand is ouder van zijn of haar over-overgrootvader/over-overgrootmoeder (gaat oneindig door)
4. Iemand is ouder van zijn of haar neef/nicht
5. Iemand is ouder van zijn of haar oom/tante
6. Iemand is ouder van zijn of haar klein-kleinkinderen
7. Iemand is ouder van zijn of haar klein-klein-kleinkinderen (gaat oneindig door)

###### Regels die de andere vreemde relaties verbieden
 | Formule | Uitleg |
 | ------- | ------ |
 | $\neg \exists mwxyz(R(x,y) \wedge R(x,z) \wedge R(x,w) \wedge R(m,z) \wedge R(m,y) \wedge R(m,w) \wedge R(y,z) \wedge R(y,w)) $| Niemand kan ouder zijn van zijn broer/zus |
 | $\neg \exists xwzy(R(y,z) \land R(z,w) \land R(w,x) \land R(y,x))$ | Niemand kan ouder zijn van zijn overgrootopa/overgrootoma |
 | $\neg \exists xwzqy(R(y,z) \land R(z,w) \land R(w,q) \land R(q,x) \land R(y,x))$ | Niemand kan ouder zijn van zijn overovergrootopa/overovergrootoma (gaat oneindig door) |
 | $\neg \exists pvwxy(R(x,y) \wedge R(w,x) \wedge R(w,v) \wedge R(v,p) \wedge R(y,p))$ | Niemand kan ouder zijn van zijn neef/nicht |
 | $\neg \exists xwzy(R(x,y) \wedge R(w,x) \wedge R(w,v) \wedge R(y,v)  )$ | Niemand kan ouder zijn van zijn oom/tante|
 | $\neg \exists xwzy(R(y,z) \land R(z,w) \land R(w,x) \land R(x,y))$ | Niemand is ouder van zijn of haar klein-kleinkinderen |
 | $\neg \exists xwzqy(R(y,z) \land R(z,w) \land R(w,q) \land R(q,x) \land R(x,y))$ | Niemand is ouder van zijn of haar klein-klein-kleinkinderen (gaat oneindig door) |


De link naar de afbeelding van de modellen:
https://ibb.co/z7Hq82Z

![Nog meer vreemde modellen](https://i.ibb.co/yYBHhkq/model-relaties.png)

# Q3 Conjunctieve vragen aan een netwerk (4pt)

* Een netwerk met feiten (en regels) is natuurlijk ook een database (of een *triple store*), en wat doen we met dat soort dingen? We stellen er vragen aan, en krijgen dan antwoord. Normaal doe je dat in SQL. In deze cursus leer je hoe je dat doet in SPARQL. Later bij Data Science leer je hoe je dat doet met `pandas`. 
* De preciese taal maakt niet, de meeste vragen die je stelt in je leven (en alle andere mensen op aarde) zijn zogenaamde *conjunctive queries*: eerste orde formules van de vorm $\exists \bar{x} (\bigwedge R(x_i,x_j))$. 
* Bijvoorbeeld, "wie zijn de grootouders van Henkie?" wordt dan  $\exists z (R(x,z) \wedge R(z,"Henkie"))$. Je kunt dat zelf vast prima in SQL en volgende week in SPARQL zeggen.
* Een *conjucntieve query* kan je ook zien als een netwerk! Als een voorbeeld netwerkje.
    * $R(x,y)\wedge R(y,z) \wedge R(x,z)$ vraagt om alle "driehoekjes"
    * Dat kan ik ook zeggen als dit netwerk (ja ik weet dat dit ongericht is, en wij alles gericht doen, maar dat laat ik aan jullie over)
<pre>
A = nx.Graph()
A.add_edges_from([('x','y'),('x','z'),('x','z')])
</pre>
* Het **antwoord** op deze vraag gegeven een netwerk `M` is dan alle deelnetwerken van `M` die isomorf zijn met `A` (plus natuurlijk hoe de knopen in `M` dan gematched worden met de "knopen" (= variabelen) in `A`.
* Programmeer dit in networkx.
* Het antwoord is verbazend simpel, uiteindelijk. Denk eerst na, ga dan gericht zoeken!

#### Eisen
* `CQ` kan werken op `nx.Digraphs` met boolse attributen op de knopen, en 1 relatie. 
* Dus vergelijkbaar met je ouder database (maar dan heb je twee boolse attributen). Je hebt dus een attribuut voor elke 1 plaatsige relatie.
* Je kan ervan uitgaan dat de query graaf altijd in de "taal" geschreven is van de model graaf.
* Let op: je hoeft de projectie niet te doen in CQ. Dus je vraag is  altijd  equivalent aan een een conjunctie van atomen (dus zonder existentiele quantifier (dat is namelijk de projectie in DB termen)).

## Deel 2

* Hoe kostbaar is deze manier van vragen beantwoorden? 
    * Probeer het eens uit te rekenen als functie van  het aantal knopen in het Model netwerk en in het vraag netwerk. 
    * Heftig?
    * Zou het beter kunnen?
    * Zou het echt beter kunnen?
    * Zoek eens op hoe lastig dit probleem eigenlijk is. Had je niet gedacht he?


```python
G= nx.karate_club_graph()
#nx.draw_networkx(G)

# Geef me alle driehoekjes in G. Ik vond 270 verschillende.
A = nx.Graph()
A.add_edges_from([('x','y'),('x','z'),('y','z')])
# A.add_edges_from([(x,y) for x in G.nodes() for y in G.nodes() for z in G.nodes() if G.has_edge(x,y) and G.has_edge(x,z) and G.has_edge(y,z)])
# nx.draw_networkx(A)

lijst_met_alle_driehoekjes = [(x,y,z) for x in G.nodes() for y in G.nodes() for z in G.nodes() if G.has_edge(x,y) and G.has_edge(x,z) and G.has_edge(y,z)]

print(len(lijst_met_alle_driehoekjes))
```

Output:
```
270
```

```python
type(A)
```

Output:
```
networkx.classes.graph.Graph
```

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def CQ(Model,Query):
    '''Returns a generator of dicts, where each dict is a bijective mapping from nodes in Model
    to all nodes in  Query.
    The keys are the model nodes.
    The set of dict values equals the set of nodes of Query.
    The generator returns all answers to the Query on this Model.'''
    if type(Model) == type(Query) == type(nx.Graph()):
        return nx.algorithms.isomorphism.GraphMatcher(Model,Query).subgraph_isomorphisms_iter()
    elif type(Model) == type(Query) == type(nx.DiGraph()):
        return nx.algorithms.isomorphism.DiGraphMatcher(Model,Query).subgraph_isomorphisms_iter()
    elif type(Model) != type(Query):
        print("Query must be of the same type as Model.")
    else:
        print("Only graphs of type Graph or DiGraph are accepted")
    
#test
# print(len(list(CQ(G,A))))
# for antwoord in CQ(G,Ag
#     print(antwoord)

CQ(G,A)
```

Output:
```
<generator object GraphMatcher.subgraph_isomorphisms_iter at 0x7ff738e4dd60>
```

```python
testG = G.subgraph([i for i in range(20)])

print(len(list(CQ(G,A))))
```

Output:
```
270
```

<!-- grade_id: q3t, points: 4, grade: true, solution: false -->
```python
assert isinstance(CQ(G,A),Generator)

```

```python

```