# Assignment Kennis en Data

## Notebook made by


|Naam|Studentnummer|
|-|-|
| | |
| | |
| | |
| | |

## Peerreview

* Vul onderstaande dict/tabel in en run de cel. Klopt het? Zijn jullie het er allemaal mee eens?
    * Verander `1,2,3,4` door jullie studentnummers
    * `1: {1:.2,2:.3,3:.3,4:.2}` betekent dat student 1 zijn punt zo over de 4 studenten verdeelt
    * Let op, per student moeten de punten optellen tot 1!
    * De gegeven punten staan in de tabel op de rijen, de ontvangen punten in de kolommen.
    * met `cijfer_per_student` kan je bepalen welk cijfer ieder krijgt bij een gegeven groepscijfer.

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.
* Elke vraag is 1 punt waard, tenzij anders aangegeven. Soms is die punt onderverdeeld in deelpunten, maar niet altijd. 

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

## Na het inleveren!

* Het gebeurt erg vaak dat mensen een "leeg bestand" inleveren. Vaak een andere versie van de opgave die nog ergens op je computer rondslingerde. Zonde van al je werk toch!
* Dus, lever **minstens een half uur voor tijd in**. Download dan wat je hebt ingeleverd op Canvas. Geef het een andere naam om verwarring te voorkomen. En draai alle cellen, en bekijk het. Geen syntax fouten? Alle vragen gemaakt? Dan zit het vast wel goed, en hoef je niet in de zenuwen te zitten.

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Q1-logisch-gevolg-(4pts)" data-toc-modified-id="Q1-logisch-gevolg-(4pts)-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Q1 logisch gevolg (4pts)</a></span></li><li><span><a href="#Q2-Vreemde-modellen-(4pts)" data-toc-modified-id="Q2-Vreemde-modellen-(4pts)-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Q2 Vreemde modellen (4pts)</a></span></li><li><span><a href="#Q3-Conjunctieve-vragen-aan-een-netwerk-(4pts)" data-toc-modified-id="Q3-Conjunctieve-vragen-aan-een-netwerk-(4pts)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Q3 Conjunctieve vragen aan een netwerk (4pts)</a></span></li><li><span><a href="#Q1-logisch-gevolg" data-toc-modified-id="Q1-logisch-gevolg-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Q1 logisch gevolg</a></span></li><li><span><a href="#Q2-Vreemde-modellen" data-toc-modified-id="Q2-Vreemde-modellen-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Q2 Vreemde modellen</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Deel-1-(2pt)" data-toc-modified-id="Deel-1-(2pt)-5.0.1"><span class="toc-item-num">5.0.1&nbsp;&nbsp;</span>Deel 1 (2pt)</a></span></li><li><span><a href="#Deel-2-(2pt)" data-toc-modified-id="Deel-2-(2pt)-5.0.2"><span class="toc-item-num">5.0.2&nbsp;&nbsp;</span>Deel 2 (2pt)</a></span></li><li><span><a href="#deel-1" data-toc-modified-id="deel-1-5.0.3"><span class="toc-item-num">5.0.3&nbsp;&nbsp;</span>deel 1</a></span></li><li><span><a href="#deel-2" data-toc-modified-id="deel-2-5.0.4"><span class="toc-item-num">5.0.4&nbsp;&nbsp;</span>deel 2</a></span></li></ul></li></ul></li><li><span><a href="#Q3-Conjunctieve-vragen-aan-een-netwerk-(4pt)" data-toc-modified-id="Q3-Conjunctieve-vragen-aan-een-netwerk-(4pt)-6"><span class="toc-item-num">6&nbsp;&nbsp;</span>Q3 Conjunctieve vragen aan een netwerk (4pt)</a></span><ul class="toc-item"><li><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Eisen" data-toc-modified-id="Eisen-6.0.0.1"><span class="toc-item-num">6.0.0.1&nbsp;&nbsp;</span>Eisen</a></span></li></ul></li></ul></li><li><span><a href="#Deel-2" data-toc-modified-id="Deel-2-6.1"><span class="toc-item-num">6.1&nbsp;&nbsp;</span>Deel 2</a></span></li></ul></li></ul></div>

```python
from typing import Generator
from nose.tools import assert_equal, assert_count_equal
import networkx as nx
%matplotlib inline
```

# Q1 logisch gevolg (4pts)

* Je maakt een Python programma waaraan je een verzameling premissen $\Sigma$ en een conclusie $\phi$ in propositielogische formules geeft en die teruggeeft of $\Sigma\models\phi$. 
    * In het geval van niet geef je een (of alle) tegenvoorbeelden terug.
    
# Q2 Vreemde modellen (4pts)

* Je maakt vreemde modellen voor je netwerk van familierelaties. Modellen die wel mogen, maar in het echt niet of vrijwel nooit voorkomen. 
* Daarna probeer je die met logische formules te verbieden. Je verzint dus zinnen die *niet waar* zijn op die ongewenste modellen.

# Q3 Conjunctieve vragen aan een netwerk (4pts)

* Een netwerk met feiten (en regels) is natuurlijk ook een database (of een *triple store*), en wat doen we met dat soort dingen? We stellen er vragen aan, en krijgen dan antwoord. Normaal doe je dat in SQL. In deze cursus leer je hoe je dat doet in SPARQL. Later bij Data Science leer je hoe je dat doet met `pandas`. 
* De preciese taal maakt niet, de meeste vragen die je stelt in je leven (en alle andere mensen op aarde) zijn zogenaamde *conjunctive queries*: eerste orde formules van de vorm $\exists \bar{x} (\bigwedge R(x_i,x_j))$. 
    * OK OK, je gebruikt ook nog wel eens het connectief "of", maar dat is makkelijk toe te voegen aanje programma als je dit hebt.
        * (want je kunt alle "of"s naar buiten brengen, en dan heb je dus een grote disjunctie van CQs.
        * die draai je allemaal achter (of naast) elkaar, en neemt de vereniging van de antwoorden.
        * Probeer maar eens een voorbeeld
* Bijvoorbeeld, "wie zijn de grootouders van Henkie?" wordt dan  $\exists z (R(x,z) \wedge R(z,"Henkie"))$. Je kunt dat zelf vast prima in SQL en volgende week in SPARQL zeggen.

# Q1 logisch gevolg

* We werken met de [sympy module](https://docs.sympy.org/latest/modules/logic.html). Dat is erg handig. Neem hem eerst goed door en speel er wat mee.
* In de cel hieronder vind je een handige routine.
* Je gebruikt natuurlijk  de definitie van $\Sigma\models\phi$, van logisch gevolg dus.
* En daarvoor is de `sympy` routine `satisfiable` best handig....
* De tests in de antwoord cel gaven bij mij dit:
<pre>
{p, q | ~p} q True
{q, q | ~p} p [{q: True, p: False}]
{q, Implies(p, q)} p {q: True, p: False}
{q & (r | s), Implies(p, q)} p [{q: True, r: True, s: True, p: False}, {q: True, r: True, s: False, p: False}, {q: True, s: True, r: False, p: False}]
WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.
</pre>

```python
from sympy import *
from sympy.logic.inference import satisfiable

p,q,r,s,t = symbols('p q r s t')

def maakconjunctie(P):
    '''Creates a conjunction from a set or list of formulas.'''
    P=list(P)
    if P:
        return P[0]&maakconjunctie(P[1:])
    else:
        return True
#test
maakconjunctie({p, ~p |q, p|(p&q)})    
```

Output:
```
p & (q | ~p) & (p | (p & q))
```

```python
C = p

Set_C = set()
Set_C.add(C)

maakconjunctie(Set_C)    
```

```python
from sympy.parsing.sympy_parser import parse_expr
import itertools

def solve_conjunction_With_Premis_and_Conclusion(Premisses,Conclusion,allmodels=False):

    TString = maakconjunctie(Premisses)
    S = str(TString)
    pol = parse_expr(S)
    values = [False, True]
    letters = list(pol.atoms())
    length = len(pol.atoms())
    permutations = list(itertools.product(values, repeat=length))
    d = dict()
    Badies = list()

    CString = maakconjunctie(Conclusion)
    CS = str(CString)
    Cpol = parse_expr(CS)
    
    for i in range(len(permutations)):
        option = permutations[i]
        for index in range(len(letters)):
            d[letters[index]] = option[index]
        if pol.subs(d)>>Cpol.subs(d) == False:
            Badies.append(d)
        
        d = dict()
        
    if Badies:
        if allmodels == True:
            return Badies
        else:
            return Badies[0]
    else:
        return True
    
    
# P={q & (r | s), Implies(p, q)}
# C= {p}
# print(solve_conjunction_3(P,C,allmodels=True))

P= {~p|q,p}
C= {q}   

print(P, C, solve_conjunction_With_Premis_and_Conclusion(P,C, allmodels=True ))
```

Output:
```
{q | ~p, p} {q} True
```

```python
from sympy.parsing.sympy_parser import parse_expr
import itertools

def solve_conjunction_2(Premisses,Conclusion,allmodels=False):

    TString = maakconjunctie(Premisses)
    S = str(TString)
    pol = parse_expr(S)
    values = [False, True]
    letters = list(pol.atoms())
    length = len(pol.atoms())
    permutations = list(itertools.product(values, repeat=length))
    d = dict()
    Badies = list()

    CString = str(Conclusion)
    Clength = len(list(CString))
#     Cpermutations = list(itertools.product(values, repeat=Clength)) 

    for i in range(len(permutations)):
        option = permutations[i]
        for index in range(len(letters)):
            d[letters[index]] = option[index]
        for f in d:
            if str(f) == CString:
                if (pol.subs(d)>>d.get(f)) == False:
                    Badies.append(d)
        d = dict()

    if Badies:
        if allmodels == True:
            return Badies
        else:
            return Badies[0]
    else:
        return True

    
    
#   {q, Implies(p, q)} p {q: True, p: False}  

P={q & (r | s), Implies(p, q)}
C= 'p'

P2 = {q, Implies(p, q)}
C2 = 'p'
# print(solve_conjunction_2({q & (r | s), Implies(p, q)},'p',allmodels=True))
print(solve_conjunction_2(P,C,allmodels=True))
print('\n')
print(solve_conjunction_2(P2,C2,allmodels=True))

```

Output:
```
[{q: True, s: False, r: True, p: False}, {q: True, s: True, r: False, p: False}, {q: True, s: True, r: True, p: False}]


[{q: True, p: False}]
```

```python
from sympy.parsing.sympy_parser import parse_expr
import itertools

def solve_conjunction(conjunction):
    TString = maakconjunctie(conjunction)
    S = str(TString)
    pol = parse_expr(S)
    values = [False, True]
    letters = list(pol.atoms())
    length = len(pol.atoms())
    permutations = list(itertools.product(values, repeat=length))
    d = dict()
    Badies = dict()
    for i in range(len(permutations)):
        option = permutations[i]
        for index in range(len(letters)):
            d[letters[index]] = option[index]
        if pol.subs(d) == False:
            Badies[pol] = d
        d = dict()
    if Badies:
        return Badies
    else:
        return True


print(solve_conjunction({p>>q,q,(r|s)}))
print(solve_conjunction({p|~p}))


```

Output:
```
{q & (Implies(p, q)) & (r | s): {q: True, s: False, r: False, p: True}}
True
```

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
from sympy.parsing.sympy_parser import parse_expr
import itertools
import numpy 

# class Error(Exception):
#     """Base class for other exceptions"""
#     pass

# class BadConclusionlError(Error):
#     "Raised not all proposition letters in the Conclusion occur in some formula of the Premisses"
#     pass

def implies(Premisses, Conclusion, allmodels=False):
    '''Return True if Conclusion follows from Premisses, 
    return a (or all) countermodel(s) otherwise. '''
    P= maakconjunctie(Premisses)
    Set_C = set()
    Set_C.add(Conclusion)
    Conclusion = maakconjunctie(Set_C)
    
    
    if not Conclusion.atoms().issubset(P.atoms()):
        print('WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.')
#         raise BadConclusionlError
    else:
        Conclusion = Set_C
        TString = maakconjunctie(Premisses)
        S = str(TString)
        pol = parse_expr(S)
        values = [False, True]
        letters = list(pol.atoms())
        length = len(pol.atoms())
        permutations = list(itertools.product(values, repeat=length))
        d = dict()
        Badies = list()

        CString = maakconjunctie(Conclusion)
        CS = str(CString)
        Cpol = parse_expr(CS)

        for i in range(len(permutations)):
            option = permutations[i]
            for index in range(len(letters)):
                d[letters[index]] = option[index]
            if pol.subs(d)>>Cpol.subs(d) == False:
                Badies.append(d)

            d = dict()

        if Badies:
            if allmodels == True:
                return Badies
            else:
                return Badies[0]
        else:
            return True
    
     
#TEST
P = {~p|q,p}
C = q  

print(P, C, implies(P,C, allmodels=True ))
P = {~p|q,q}
C = p 

print(P, C, implies(P,C, allmodels=True ))

P = {p>>q,q}  # andere manier om implicatie te schrijven
C = p   

print(P, C, implies(P,C))

P = {p>>q,q,(r|s)}  # andere manier om implicatie te schrijven
C = p

print(P, C, implies(P,C, allmodels=True ))
P = {p>>q,q}  # andere manier om implicatie te schrijven
C = r

print(implies(P,C,allmodels=False))
#implies(P,C) Dit geeft een error. Probeer het ook zelf

# {p, q | ~p} q True
# {q, q | ~p} p [{q: True, p: False}]
# {q, Implies(p, q)} p {q: True, p: False}
# {q & (r | s), Implies(p, q)} p [{q: True, r: True, s: True, p: False}, {q: True, r: True, s: False, p: False}, {q: True, s: True, r: False, p: False}]
# WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.
```

Output:
```
{q | ~p, p} q True
{q, q | ~p} p [{q: True, p: False}]
{q, Implies(p, q)} p {q: True, p: False}
{q, Implies(p, q), r | s} p [{q: True, s: False, r: True, p: False}, {q: True, s: True, r: False, p: False}, {q: True, s: True, r: True, p: False}]
WARNING: All proposition letters in the Conclusion should occur in some formula of the Premisses.
None
```

```python
maakconjunctie({q & (r | s), Implies(p, q)}) 
# {q, Implies(p, q), r | s}  
```

Output:
```
q & (Implies(p, q)) & (r | s)
```

```python
maakconjunctie({q, Implies(p, q), r | s} ) 
```

Output:
```
q & (Implies(p, q)) & (r | s)
```

<!-- grade_id: q1t, points: 4, grade: true, solution: false -->
```python
assert_equal(type(implies({p},p)), bool)
assert_equal(type(implies({p|q},p)), dict)
assert_equal(type(implies({p|q},p, allmodels=True)), list)

```

# Q2 Vreemde modellen

### Deel 1 (2pt)
1. Je maakt vreemde modellen voor je netwerk van familierelaties. Modellen die wel mogen, maar in het echt niet of vrijwel nooit voorkomen. 
2. Daarna probeer je die met logische formules te verbieden. Je verzint dus zinnen die *niet waar* zijn op die ongewenste modellen.

In ieder geval verzin je zinnen die de volgende vijf modellen verbieden (die soort $\Sigma$ bij een knoop is eigenlijk een M en staat voor Man. Gebruik dus weer de predicaatlogische relaties $R,M,V$ in die zinnen, en niks anders. Geef duidelijk aan welk model je met welke zin verbiedt.

![](vreemd.jpg)

### Deel 2 (2pt)

Ga nu verder en verzin nog meer modellen die je zou willen verbieden. Misschien vindt je er wel oneindig veel! Kan je daar een zin die die allemaal verbiedt voor vinden? Kan je het met oneindig veel zinnen? Geef ze!

**Je antwoorden zijn keurige Latex formules, en natuurlijk wel gevormd, met nette haakjes etc.**

**Voor je modellen, maak je plaatjes in netwerkx, met de hand, of iets anders, als het maar heel duidelijk is (en dus niet zo als die in deze opgave ;-)**

<!-- grade_id: q2a, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

# Q3 Conjunctieve vragen aan een netwerk (4pt)

* Een netwerk met feiten (en regels) is natuurlijk ook een database (of een *triple store*), en wat doen we met dat soort dingen? We stellen er vragen aan, en krijgen dan antwoord. Normaal doe je dat in SQL. In deze cursus leer je hoe je dat doet in SPARQL. Later bij Data Science leer je hoe je dat doet met `pandas`. 
* De preciese taal maakt niet, de meeste vragen die je stelt in je leven (en alle andere mensen op aarde) zijn zogenaamde *conjunctive queries*: eerste orde formules van de vorm $\exists \bar{x} (\bigwedge R(x_i,x_j))$. 
* Bijvoorbeeld, "wie zijn de grootouders van Henkie?" wordt dan  $\exists z (R(x,z) \wedge R(z,"Henkie"))$. Je kunt dat zelf vast prima in SQL en volgende week in SPARQL zeggen.
* Een *conjucntieve query* kan je ook zien als een netwerk! Als een voorbeeld netwerkje.
    * $R(x,y)\wedge R(y,z) \wedge R(x,z)$ vraagt om alle "driehoekjes"
    * Dat kan ik ook zeggen als dit netwerk (ja ik weet dat dit ongericht is, en wij alles gericht doen, maar dat laat ik aan jullie over)
<pre>
A = nx.Graph()
A.add_edges_from([('x','y'),('x','z'),('x','z')])
</pre>
* Het **antwoord** op deze vraag gegeven een netwerk `M` is dan alle deelnetwerken van `M` die isomorf zijn met `A` (plus natuurlijk hoe de knopen in `M` dan gematched worden met de "knopen" (= variabelen) in `A`.
* Programmeer dit in networkx.
* Het antwoord is verbazend simpel, uiteindelijk. Denk eerst na, ga dan gericht zoeken!

#### Eisen
* `CQ` kan werken op `nx.Digraphs` met boolse attributen op de knopen, en 1 relatie. 
* Dus vergelijkbaar met je ouder database (maar dan heb je twee boolse attributen). Je hebt dus een attribuut voor elke 1 plaatsige relatie.
* Je kan ervan uitgaan dat de query graaf altijd in de "taal" geschreven is van de model graaf.
* Let op: je hoeft de projectie niet te doen in CQ. Dus je vraag is  altijd  equivalent aan een een conjunctie van atomen (dus zonder existentiele quantifier (dat is namelijk de projectie in DB termen)).

## Deel 2

* Hoe kostbaar is deze manier van vragen beantwoorden? 
    * Probeer het eens uit te rekenen als functie van  het aantal knopen in het Model netwerk en in het vraag netwerk. 
    * Heftig?
    * Zou het beter kunnen?
    * Zou het echt beter kunnen?
    * Zoek eens op hoe lastig dit probleem eigenlijk is. Had je niet gedacht he?


```python
G= nx.karate_club_graph()
nx.draw_networkx(G)

# Geef me alle driehoekjes in G. Ik vond 270 verschillende.
A = nx.Graph()
A.add_edges_from([('x','y'),('x','z'),('y','z')])
```

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def CQ(Model,Query):
    '''Returns a generator of dicts, where each dict is a bijective mapping from  nodes in Model
    to all nodes in  Query.
    The keys are the model nodes.
    The set of dict values equals the set of nodes of Query.
    The generator returns all answers to the Query on this Model.'''
    # YOUR CODE HERE
    raise NotImplementedError()
    
#test
print(len(list(CQ(G,A))))
#for antwoord in CQ(G,A):
#    print(antwoord)

    
```

<!-- grade_id: q3t, points: 4, grade: true, solution: false -->
```python
assert isinstance(CQ(G,A),Generator)

```